Individuality across environmental context in Drosophila melanogaster

Author response:

The following is the authors’ response to the original reviews

Public Reviews:

Reviewer #1 (Public Review):

Summary:

The authors state the study's goal clearly: "The goal of our study was to understand to what extent animal individuality is influenced by situational changes in the environment, i.e., how much of an animal's individuality remains after one or more environmental features change." They use visually guided behavioral features to examine the extent of correlation over time and in a variety of contexts. They develop new behavioral instrumentation and software to measure behavior in Buridan's paradigm (and variations thereof), the Y-maze, and a flight simulator. Using these assays, they examine the correlations between conditions for a panel of locomotion parameters. They propose that inter-assay correlations will determine the persistence of locomotion individuality.

Strengths:

The OED defines individuality as "the sum of the attributes which distinguish a person or thing from others of the same kind," a definition mirrored by other dictionaries and the scientific literature on the topic. The concept of behavioral individuality can be characterized as:

(1) a large set of behavioral attributes,

(2) with inter-individual variability, that are

(3) stable over time.

A previous study examined walking parameters in Buridan's paradigm, finding that several parameters were variable between individuals, and that these showed stability over separate days and up to 4 weeks (DOI: 10.1126/science.aaw718). The present study replicates some of those findings and extends the experiments from temporal stability to examining the correlation of locomotion features between different contexts.

The major strength of the study is using a range of different behavioral assays to examine the correlations of several different behavior parameters. It shows clearly that the inter-individual variability of some parameters is at least partially preserved between some contexts, and not preserved between others. The development of highthroughput behavior assays and sharing the information on how to make the assays is a commendable contribution.

We thank the reviewer for his exceptionally kind assessment of our work!

Weaknesses:

The definition of individuality considers a comprehensive or large set of attributes, but the authors consider only a handful. In Supplemental Fig. S8, the authors show a large correlation matrix of many behavioral parameters, but these are illegible and are only mentioned briefly in Results.

We have now uploaded a high-resolution PDF to the Github Address: https://github.com/LinneweberLab/Mathejczyk_2024_eLife_Individuality/blob/main/S8.pdf, and this is also mentioned in the figure legend for Fig. S8

Why were five or so parameters selected from the full set? How were these selected?

The five parameters (% of time walked, walking speed, vector strength, angular velocity, and centrophobicity) were selected because they describe key aspects of the investigated behaviors that can be compared directly across assays. Importantly, several parameters we typically use (e.g., Linneweber et al., 2020) cannot be applied under certain conditions, such as darkness or the absence of visual cues. Furthermore, these five parameters encompass three critical aspects of navigation across standard visual behavioral arenas: (1) The “exploration” category is characterized by parameters describing the fly’s activity. (2) Parameters related to “attention” reflect heightened responses to visual cues, but unlike commonly used metrics such as angle or stripe deviations (e.g., Coulomb, 2012; Linneweber et al., 2020), they can also be measured in absence of visual cues and are therefore suitable for cross-assay comparisons. (3) The parameter “centrophobicity,” used as a potential indicator of anxiety, is conceptually linked to the open-field test in mice, where the ratio of wall-to-open-field activity is frequently calculated as a measurement of anxiety (see for example Carter, Sheh, 2015, chapter 2. https://www.sciencedirect.com/book/9780128005118/guide-to-researchtechniques-in-neuroscience). Admittedly, this view is frequently challenged in mice, but it has a long history which is why we use it.

Do the correlation trends hold true across all parameters? For assays in which only a subset of parameters can be directly compared, were all of these included in the analysis, or only a subset?

As noted above, we only included a subset of parameters in our final analysis, as many were unsuitable for comparison across assays while still providing valuable assayspecific information which are important to relate these results to previous publications.

The correlation analysis is used to establish stability between assays. For temporal retesting, "stability" is certainly the appropriate word, but between contexts, it implies that there could be 'instability'. Rather, instead of the 'instability' of a single brain process, a different behavior in a different context could arise from engaging largely (or entirely?) distinct context-dependent internal processes, and have nothing to do with process stability per se. For inter-context similarities, perhaps a better word would be "consistency".

Thank you for this suggestion. During the preparation of the manuscript, we indeed frequently alternated between the terms “stability” and “consistency.” And decided to go with “stability” as the only descriptor, to keep it simple. We now fully agree with the reviewer’s argument and have replaced “stability” by “consistency” throughout the current version of the manuscript in order to increase clarity and coherence.

The parameters are considered one by one, not in aggregate. This focuses on the stability/consistency of the variability of a single parameter at a time, rather than holistic individuality. It would appear that an appropriate measure of individuality stability (or individuality consistency) that accounts for the high-dimensional nature of individuality would somehow summarize correlations across all parameters. Why was a multivariate approach (e.g. multiple regression/correlation) not used? Treating the data with a multivariate or averaged approach would allow the authors to directly address 'individuality stability' and analyses of single-parameter variability stability.

We agree with the reviewer that a multivariate analysis adds clear advantages in terms of statistical power, in addition to our chosen approach. On one hand, we believe that the simplicity of our initial analysis, both for correlational and mean data, makes easy for readers to understand and reproduce our data. While preparing the previous version of the manuscript we were skeptical since more complex analyses often involve numerous choices, which can complicate reproducibility. For instance, a recent study in personality psychology (Paul et al., 2024) highlighted the risks of “forking paths” in statistical analysis, showing that certain choices of statistical methods could even reverse findings—a concern mitigated by our simplistic straightforward approach. Still, in preparation of this revised version of the manuscript, we accepted the reviewer’s advice and reanalyzed the data using a generalized linear model. This analysis nicely recapitulates our initial findings and is now summarized in a single figure (Fig. 9).

The correlation coefficients are sometimes quite low, though highly significant, and are deemed to indicate stability. For example, in Figure 4C top left, the % of time walked at 23{degree sign}C and 32{degree sign}C are correlated by 0.263, which corresponds to an R2 of 0.069 i.e. just 7% of the 32{degree sign}C variance is predictable by the 23{degree sign}C variance. Is it fair to say that a 7% determination indicates parameter stability? Another example: "Vector strength was the most correlated attention parameter... correlations ranged... to -0.197," which implies that 96% (1 - R2) of Y-maze variance is not predicted by Buridan variance. At what level does an r value not represent stability?

We agree that this is an important question. Our paper clearly demonstrates that individuality always plays a role in decision-making (and, in this context, any behavioral output can be considered a decision). However, the non-linear relationship between certain situations and the individual’s behavior often reduces the predictive value (or correlation) across contexts, sometimes quite drastically.

For instance, temperature has a relatively linear effect on certain behavioral parameters, leading to predictable changes across individuals. As a result, correlations across temperature conditions are often similar to those observed across time within the same situation. In contrast, this predictability diminishes when comparing conditions like the presence or absence of visual stimuli, the use of different arenas, or different modalities.

For this reason, we believe that significance remains the best indicator for describing how measurable individuality persists, even across vastly different situations.

The authors describe a dissociation between inter-group differences and interindividual variation stability, i.e. sometimes large mean differences between contexts, but significant correlation between individual test and retest data. Given that correlation is sensitive to slope, this might be expected to underestimate the variability stability (or consistency). Is there a way to adjust for the group differences before examining the correlation? For example, would it be possible to transform the values to in-group ranks prior to correlation analysis?

We thank the reviewer for this suggestion, and we have now addressed this point. To account for slope effects, we have now introduced in-group ranks for our linear model computation (see Fig. 9).

What is gained by classifying the five parameters into exploration, attention, and anxiety? To what extent have these classifications been validated, both in general and with regard to these specific parameters? Is the increased walking speed at higher temperatures necessarily due to an increased 'explorative' nature, or could it be attributed to increased metabolism, dehydration stress, or a heat-pain response? To what extent are these categories subjective?

We agree that grouping our parameters into traits like exploration, attention, and anxiety always includes subjective decisions. The classification into these three categories is even considered partially controversial in the mouse specific literature, which uses the term “anxiety” in similar experiments (see for exampler Carter, Sheh, 2015, chapter 2 . https://www.sciencedirect.com/book/9780128005118/guide-to-research-techniquesin-neuroscience). Nevertheless, we believe that readers greatly benefit from these categories, since they make it easier to understand (beyond mathematical correlations) which aspects of the flies’ individuality can be considered consistent across situations. Furthermore, these categories serve as a bridge to compare insight from very distinct models.

The legends are quite brief and do not link to descriptions of specific experiments. For example, Figure 4a depicts a graphical overview of the procedure, but I could not find a detailed description of this experiment's protocol.

We assume the reviewer is referring to Figure 3a. The detailed experimental protocol can be found in the Materials and Methods section under Setup 2: IndyTrax Multi-Arena Platform. We have now clarified this in the mentioned figure legend.

Using the current single-correlation analysis approach, the aims would benefit from rewording to appropriately address single-parameter variability stability/consistency (as distinct from holistic individuality). Alternatively, the analysis could be adjusted to address the multivariate nature of individuality, so that the claims and the analysis are in concordance with each other.

The reviewer raises an important point about hierarchies within the concept of animal individuality or personality. We agree that this is best addressed by first focusing on single behavioral traits/parameters and then integrating several trait properties into a cohesive concept of animal personality (holistic individuality). To ensure consistency throughout the text, we have now thoroughly reviewed the entire manuscript clearly distinguish between single-parameter variability stability/consistency and holistic individuality/personality.

The study presents a bounty of new technology to study visually guided behaviors. The GitHub link to the software was not available. To verify the successful transfer of open hardware and open-software, a report would demonstrate transfer by collaboration with one or more other laboratories, which the present manuscript does not appear to do. Nevertheless, making the technology available to readers is commendable.

We have now uploaded all codes and materials to GitHub and made them available as soon as we received the reviewers’ comments. All files and materials can be accessed at https://github.com/LinneweberLab/Mathejczyk_2024_eLife_Individuality, which is now frequently mentioned throughout the revised manuscript.

The study discusses a number of interesting, stimulating ideas about inter-individual variability, and presents intriguing data that speaks to those ideas, albeit with the issues outlined above.

While the current work does not present any mechanistic analysis of inter-individual variability, the implementation of high-throughput assays sets up the field to more systematically investigate fly visual behaviors, their variability, and their underlying mechanisms.

We thank the reviewer again for the extensive and constructive feedback.

Reviewer #2 (Public Review):

Summary:

The authors repeatedly measured the behavior of individual flies across several environmental situations in custom-made behavioral phenotyping rigs.

Strengths:

The study uses several different behavioral phenotyping devices to quantify individual behavior in a number of different situations and over time. It seems to be a very impressive amount of data. The authors also make all their behavioral phenotyping rig design and tracking software available, which I think is great and I'm sure other folks will be interested in using and adapting it to their own needs.

We thank the reviewer for highlighting the strengths of our study.

Weaknesses/Limitations:

I think an important limitation is that while the authors measured the flies under different environmental scenarios (i.e. with different lighting and temperature) they didn't really alter the "context" of the environment. At least within behavioral ecology, context would refer to the potential functionality of the expressed behaviors so for example, an anti-predator context, a mating context, or foraging. Here, the authors seem to really just be measuring aspects of locomotion under benign (relatively low-risk perception) contexts. This is not a flaw of the study, but rather a limitation to how strongly the authors can really say that this demonstrates that individuality is generalized across many different contexts. It's quite possible that rank order of locomotor (or other) behaviors may shift when the flies are in a mating or risky context.

We agree with the reviewer that the definition of environmental context can differ between fields and that behavioral context is differently defined, particularly in ecology. Nevertheless, we highlight that our alternations of environmental context are highly stereotypic, well-defined, and unbiased from any interpretation (we only modified what we stated in the experimental description without designing a specific situation that might be again perceived individually differently. E.g., comparing a context with a predator and one without might result in a binary response because one fraction of the tested individuals might perceive the predator in the predator situation, and the other half does not.

The analytical framework in terms of statistical methods is lacking. It appears as though the authors used correlations across time/situations to estimate individual variation; however, far more sophisticated and elegant methods exist. The paper would be a lot stronger, and my guess is, much more streamlined if the authors employ hierarchical mixed models to analyse these data these models could capture and estimate differences in individual behavior across time and situations simultaneously. Along with this, it's currently unclear whether and how any statistical inference was performed. Right now, it appears as though any results describing how individuality changes across situations are largely descriptive (i.e. a visual comparison of the strengths of the correlation coefficients?).

The reviewer raises an important point, also raised by reviewer #1. On one hand, we agree with both reviewers that a more aggregated analysis has clear advantages like more statistical power and has the potential to streamline our manuscript, which is why we added such an analysis (see below). On the other hand, we would also like to defend the initial approach we took, since we think that the simplicity of the analysis for both correlational and mean data is easy to understand and reproduce. More complex analyses necessarily include the selection of a specific statistical toolbox by the experimenters and based on these decisions, different analyses become less comparable and more and more complicated to reproduce, unless the entire decision tree is flawlessly documented. For instance, a recent personality psychology paper investigated the relationship between statistical paths within the decision tree (forking analysis) and their results, leading to very surprising results (Paul et al., 2024), since some paths even reversed their findings. Such a variance in conclusions is hardly possible with the rather simplistic and easily reproducible analysis we performed. One of the major strengths of our study is the simple experimental design, allowing for rather simple and easy to understand analyses.

We nevertheless took the reviewer’s advice very seriously and reanalyzed the data using a generalized linear model, which largely recapitulated the findings of our previously performed “low-tech” analysis in a single figure (Fig. 9).

Another pretty major weakness is that right now, I can't find any explicit mention of how many flies were used and whether they were re-used across situations. Some sort of overall schematic showing exactly how many measurements were made in which rigs and with which flies would be very beneficial.

We apologize for this inconvenience. A detailed overview of male and female sample sizes has been listed in the supplemental boxplots next to the plots (e.g, Fig S6). Apparently, this was not visible enough. Therefore, we have now also uniformly added the sample sizes to the main figure legends.

I don't necessarily doubt the robustness of the results and my guess is that the author's interpretations would remain the same, but a more appropriate modeling framework could certainly improve their statistical inference and likely highlight some other cool patterns as these methods could better estimate stability and covariance in individual intercepts (and potentially slopes) across time and situation.

As described above, we have now added the suggested analyses. We hope that the reviewer will appreciate the new Fig. 9, which, in our opinion, largely confirms our previous findings using a more appropriate generalized linear modelling framework.

Reviewer #3 (Public Review):

This manuscript is a continuation of past work by the last author where they looked at stochasticity in developmental processes leading to inter-individual behavioural differences. In that work, the focus was on a specific behaviour under specific conditions while probing the neural basis of the variability. In this work, the authors set out to describe in detail how stable the individuality of animal behaviours is in the context of various external and internal influences. They identify a few behaviours to monitor (read outs of attention, exploration, and 'anxiety'); some external stimuli (temperature, contrast, nature of visual cues, and spatial environment); and two internal states (walking and flying).

They then use high-throughput behavioural arenas - most of which they have built and made plans available for others to replicate - to quantify and compare combinations of these behaviours, stimuli, and internal states. This detailed analysis reveals that:

(1) Many individualistic behaviours remain stable over the course of many days.

(2) That some of these (walking speed) remain stable over changing visual cues. Others (walking speed and centrophobicity) remain stable at different temperatures.

(3) All the behaviours they tested failed to remain stable over the spatially varying environment (arena shape).

(4) Only angular velocity (a readout of attention) remains stable across varying internal states (walking and flying).

Thus, the authors conclude that there is a hierarchy in the influence of external stimuli and internal states on the stability of individual behaviours.

The manuscript is a technical feat with the authors having built many new highthroughput assays. The number of animals is large and many variables have been tested - different types of behavioural paradigms, flying vs walking, varying visual stimuli, and different temperatures among others.

We thank the reviewer for this extraordinary kind assessment of our work!

Recommendations for the authors:

Reviewing Editor (Recommendations For The Authors):

While appreciating the effort and quality of the work that went into this manuscript, the reviewers identified a few key points that would greatly benefit this work.

(1) Statistical methods adopted. The dataset produced through this work is large, with multiple conditions and comparisons that can be made to infer parameters that both define and affect the individualistic behaviour of an animal. Hierarchical mixed models would be a more appropriate approach to handle such datasets and infer statistically the influence of different parameters on behaviours. We recommend that the authors take this approach in the analyses of their data.

(2) Brevity in the text. We urge the authors to take advantage of eLife's flexible template and take care to elaborate on the text in the results section, the methods adopted, the legends, and the guides to the legends embedded in the main text. The findings are likely to be of interest to a broad audience, and the writing currently targets the specialist.

Reviewer #2 (Recommendations For The Authors):

I want to start by saying this seems like a really cool study! It's an impressive amount of work and addressing a pretty basic question that is interesting (at least I think so!)

We thank the reviewer again for this assessment!

That said, I would really strongly recommend the authors embrace using mixed/hierarchical models to analyze their data. They're producing some really impressive data and just doing Pearson correlation coefficients across time points and situations is very clunky and actually losing out on a lot of information. The most up-todate, state-of-the-art are mixed models - these models can handle very complex (or not so complex) random structures which can estimate variance and importantly, covariance, in individual intercepts both over time and across situations. I actually think this could add some really cool insights into the data and allow you to characterize the patterns you're seeing in far more detail. It's datasets exactly like this that are tailormade for these complex variance partitioning models!

As mentioned before, we have now adopted a more appropriate GLM-based data analysis (see above).

Regardless of which statistical methods you decide to use, please explicitly state in your methods exactly what analyses you did. That is completely lacking now and was a bit frustrating. As such, it's completely unclear whether or how statistical inference was performed. How did you do the behavioral clustering?

We apologize that these points were not clearly represented in the previous version of the manuscript. We have now significantly extended the methods section to include a separate paragraph on the statistical methods used, in order to address this critique and hope that the revised version is clear now.

Also, I could not for the life of me figure out how many flies had been measured. Were they reused across the situation? Or not?

We reused the same flies across situations whenever possible. However, having one fly experience all assays consecutively was not feasible due to their fragility. Instead, individual flies were exposed to at least 2 of the 3 groups of assays used here: in the Indytrax setup , the Buridan arenas and variants thereof, and the virtual arenas Hence, we have compared flies across entirely different setups, but the number of times flies can be retested is limited (as otherwise, sample sizes will drop over time, and the flies will have gone through too many experimental alternations). To make this more clear, we have elaborated on this point in the main text, and we added group sample sizes to figure legends r.

What are these "groups" and "populations" that are referred to in the results (e.g. lines 384, 391, 409)?

We apologize for using these two terms somewhat interchangeably without proper introduction/distinction. We have now made this more clear in at the beginning of the results in the main text, by focusing on the term ‘group’. By ‘group’ we refer to the average of all individuals tested in the same situation. Sample sizes in the figure legends now indicate group/population sizes to make this clearer.

Some of the rationale for the development of the behavioral rigs would have actually been nice to include in the intro, rather than in the results.

This rationale is introduced at the beginning of the last paragraph of the introduction. We hope that this now becomes clear in the revised version of the manuscript.

Reviewer #3 (Recommendations For The Authors):

This manuscript would do well to take advantage of eLife's flexible word limit. I sense that it has been written in brevity for a different journal but I would urge the authors to revisit this and unpack the language here - in the text, in the figure legends, in references to the figures within the text. The way it's currently written, though not misleading, will only speak to the super-specialist or the super-invested :). But the findings are nice, and it would be nice to tailor it to a broader audience.

We appreciate this suggestion. Initially, we were hoping that we had described our results as clearly and brief as possible. We apologize if that was not always the case. The comments and requests of all three reviewers now led to a series of additions to both main text and methods, leading to a significantly expanded manuscript. We hope that these additons are helpful for the general, non-expert audience.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Over the past decade, several studies have demonstrated that idiosyncratic animal behaviors remain stable over long time periods. The stability of individually variable behaviors over time is often referred to as an animal’s individuality, or personality. However, most experimental studies have focused on individuality in a single, well-defined environmental context, whereas it is well-established from population studies that animal behavior is highly context-dependent. The ‘person-situation debate’ in humans and decades of observations of animal individuality under intrinsically variable natural conditions raise the question of whether and to what extent animal behavior remains stable across different situations, such as changing environmental contexts. For instance, one individual might be generally more visually guided than another, or rely only on one particular visual cue, or even on this very cue only in a specific environmental context. Here, we use a combination of both well-established and novel behavioral assays to demonstrate the relationship between individual behavior and variable environmental context under tightly controlled laboratory conditions in the model system Drosophila melanogaster. The stability of three individual traits (termed exploration, attention, and anxiety) was investigated under changing environmental contexts (temperature, visual cues, arena shape), in both walking and flying flies. We find that individuality is highly context-dependent, but even under the most extreme environmental alterations tested, stability of behavior always persisted in at least one of the traits. Furthermore, our quantification reveals a hierarchical order of environmental features influencing individuality. In summary, our work demonstrates that, similar to humans, fly individuality persists across different contexts, and individual differences shape behavior across variable environments, thereby making the underlying developmental and functional mechanisms amenable to genetic dissection.

Revised version: Over the past decade, several studies have demonstrated that idiosyncratic animal behaviors remain consistent over long time periods. The consistency of individually variable behaviors over time is often referred to as an animal’s individuality, or if consisting of multiple traits personality. However, most experimental studies have focused on individuality in a single, well-defined environmental context, whereas it is well-established from population studies that animal behavior is highly context-dependent. The ‘person-situation debate’ in humans and decades of observations of animal individuality under intrinsically variable natural conditions raise the question of whether and to what extent animal behavior remains stable across different situations, such as changing environmental contexts. For instance, one individual might be generally more visually guided than another, or rely only on one particular visual cue, or even on this very cue only in a specific environmental context. Here, we use a combination of both well-established and novel behavioral assays to demonstrate the relationship between individual behavior and variable environmental context under tightly controlled laboratory conditions in the model system Drosophila melanogaster. The consistency of three individual traits (termed exploration, attention, and anxiety) was investigated under changing environmental contexts (temperature, visual cues, arena shape), in both walking and flying flies. We find that individuality is highly context-dependent, but even under the most extreme environmental alterations tested, consistency of behavioral individuality always persisted in at least one of the traits. Furthermore, our quantification reveals a hierarchical order of environmental features influencing individuality. We confirmed this hierarchy using a generalized linear model. In summary, our work demonstrates that, similar to humans, fly individuality persists across different contexts, and individual differences shape behavior across variable environments, thereby making the underlying developmental and functional mechanisms amenable to genetic dissection.

Introduction

Earlier version: Animal species display enormous interindividual differences in behavior. The temporal stability of these interindividual differences allows for defining these idiosyncrasies as animal individuality, also called animal temperament, behavioral syndrome, or even personality 1–4. Animal individuality has been described for numerous species and behaviors, including: mouse exploration behavior 5, great tit feeding preferences 6, octopus threat and feeding behavior 7, pea aphid threat responses 8, as well as vinegar fly handedness 9,10, olfaction 11, phototaxis 12, and object orientation 13. The temporal stability of individuality allows for discriminating it from short-term behavioral changes, like internal state variations. Furthermore, temporal stability of behavior allows modifications by learning 14 and behavior-based natural selection 15. Historically, the ‘nature vs. nurture’ debate (i.e., genome vs. environment) has shaped our understanding of the origins of behavioral individuality 16. Stochastic developmental processes are a more recently found factor 13,17, providing an additional layer to gene-environmental interactions 18. Most experimental studies on animal individuality have focused on a single behavior set in a singular environment, yet commonly used definitions of animal individuality in predominantly field-based disciplines such as behavioral ecology 3, and the definition of human personality entails behavioral stability across situations 19. Similar to humans, animals will never re-encounter the exact same situation under natural conditions, and individuality can be an important factor in determining the ecological success of an animal and, as a consequence, its entire population 20. It is, therefore, crucial to understand how an animal’s individuality interacts with variable environments. Despite decades of interest in animal individuality across situations, the experimental evidence remains insufficient, and the results are contradictory, with studies arguing for 21–24 and against 25,26 individuality across situations with many finding mixed results 27–29. Here we sought to systematically investigate the stability of individual behavioral differences across situations using laboratory-controlled variable environmental conditions.

Revised version: Animal species display enormous interindividual differences in behavior. The temporal consistency of these interindividual differences allows for defining these idiosyncrasies as animal individuality, also called animal temperament, behavioral syndrome, or if the entirety of trait differences between individuals is meant even animal personality (Andrew Sih et al., 2004; Dall et al., 2004; Gosling, 2008; Reale et al., 2007). Animal individuality has been described for numerous species and behaviors, including: mouse exploration behavior (Freund et al., 2013), great tit feeding preferences (Partridge, 1976), octopus threat responses and feeding behavior (Mather & Anderson, 1993), pea aphid threat responses (Schuett et al., 2011), as well as vinegar fly handedness (Ayroles et al., 2015; Buchanan et al., 2015), olfaction (Honegger et al., 2020), phototaxis (Kain et al., 2012), and object orientation (Linneweber et al., 2020). The temporal consistency of individuality allows for discriminating it from short-term behavioral changes, like internal state variations. Furthermore, temporal consistency of behavior allows modifications by learning (Kahsai & Zars, 2011) and behavior-based natural selection (Lapiedra et al., 2018). Historically, the ‘nature vs. nurture’ debate (i.e., genome vs. environment) has shaped our understanding of the origins of behavioral individuality (Galton, 1875). Stochastic developmental processes are a more recently found factor (Honegger & de Bivort, 2018; Linneweber et al., 2020), providing an additional layer to gene-environmental interactions (Traynor & Singleton, 2010). Most experimental studies on animal individuality have focused on a single behavior set in a singular environment, yet commonly used definitions of animal individuality in predominantly field-based disciplines such as behavioral ecology (Reale et al., 2007), and the definition of human personality entails behavioral consistency across situations (Mischel, 1968). Similar to humans, animals will never re-encounter the exact same situation under natural conditions, and animal personality can be an important factor in determining the ecological success of an animal and, as a consequence, its entire population (Takola & Schielzeth, 2022). It is, therefore, crucial to understand how an animal’s personality or individuality in specific traits interacts with variable environments. Despite decades of interest in animal individuality across situations, the experimental evidence remains insufficient, and the results are contradictory, with studies arguing for (Briffa et al., 2008; Herborn et al., 2010; Kralj-Fišer et al., 2007; Mowles et al., 2012) and against (Coleman & Wilson, 1998; D’Eath & Burn, 2002) individuality in animals across situations with many finding mixed results (Boissy & Bouissou, 1995; Spoolder et al., 1996; Werkhoven et al., 2021). Here we sought to systematically investigate the consistency of individual behavioral differences across situations using laboratory-controlled variable environmental conditions.

Earlier version: There is ample evidence for the stability of human personality and animal individuality over time 30,31. Intuitively, humans attribute stable personalities to other humans 32. The stability of human behavior across situations, including environmental context, nevertheless remains highly debated, with several studies arguing for 30,33 and against it 19,34. Summarized as the ’person-situation debate’, psychologists have argued for over a century about the relationship between personality and the situational environment, placing this question at the core of personality psychology 19. Animal models for behavioral stability across situations provide an opportunity to address this question. In our study, the situation is the sum of all environmental features where a specific behavioral task is performed. The work presented here establishes a quantitative animal model to investigate this and related questions about behavioral stability across environmental situations.

Revised version: There is ample evidence for the consistency of human personality and animal individuality over time (Costa & McCrae, 1988; Weisbuch et al., 2010). Intuitively, humans attribute consistent personalities to other humans (Ross, 1977). The consistency of human behavior across situations, including environmental context, nevertheless remains highly debated, with several studies arguing for (Costa & McCrae, 1988; Fleeson, 2001) and against it (Hartshorne & May, 1930; Mischel, 1968). Summarized as the ‘person-situation debate’, psychologists have argued for over a century about the relationship between personality and the situational environment, placing this question at the core of personality psychology (Mischel, 1968). Animal models for behavioral consistency across situations provide an opportunity to address this question. In our study, the situation is the sum of all environmental features where a specific behavioral task is performed. The work presented here establishes a quantitative animal model to investigate this and related questions about behavioral consistency across environmental situations.

Earlier version: The vinegar fly Drosophila melanogaster provides an enormous toolset for studying individuality and its origins across multiple scales, from neurons to circuits, to behavior. This comprises a rapidly unfolding connectome 35,36, an extensive collection of cell type-specific, binary expression systems 37,38, and a similarly extensive collection of tools for manipulating neuronal function, including neuronal silencers 39,40, activators 41, and thermogenetic 42 and optogenetic 43,44 effectors. Hence, the availability of tools and reagents for manipulating small, defined sets of neurons is no longer a major limitation in answering questions about individual behavioral differences. Instead, what has been missing is an integrative approach combining different high-throughput behavioral assays, which would allow for behavioral screening and classification of large quantities of individual flies for several different behavioral parameters across a variety of environmental contexts.

Revised version: The vinegar fly Drosophila melanogaster provides an enormous toolset for studying individuality and its origins across multiple scales, from neurons to circuits, to behavior. This comprises a rapidly unfolding connectome (Dorkenwald et al., 2024; Xu et al., 2020; Zheng et al., 2018), an extensive collection of cell type-specific, binary expression systems (Luan, Peabody, et al., 2006; Pfeiffer et al., 2010), and a similarly extensive collection of tools for manipulating neuronal function, including neuronal silencers (Baines et al., 2001; Sweeney et al., 1995), activators (Luan, Lemon, et al., 2006), and thermogenetic (Hamada et al., 2008) and optogenetic (Klapoetke et al., 2014; Pulver et al., 2009) effectors. Hence, the availability of tools and reagents for manipulating small, defined sets of neurons is no longer a major limitation in answering questions about individual behavioral differences. Instead, what has been missing is an integrative approach combining different high-throughput behavioral assays, which would allow for behavioral screening and classification of large quantities of individual flies for several different behavioral parameters across a variety of environmental contexts.

Earlier version: In recent years, mainly low-throughput, single-fly visual assays like the classical two-stripe Buridan assay 45,46, arenas for tethered walking 47 or virtual flight 48 have been used to study behavioral individuality in vison-based tasks 13,49. In Buridan’s paradigm, flies walk between two inaccessible visual targets (two vertical high-contrast stripes). The stripes are presented in a homogenously illuminated surrounding with no other visual or non-visual cues. Apart from work on individuality involving Dorsal Cluster Neurons (DCNs) 50, also known as Lobula Columnar Neuron 14 (LC14) 51, Buridan’s paradigm has been used to show the importance of the central complex for normal walking speed 52 and for studying photoreceptor function or development 53,54, and visual memory 55. Other circular arenas have been frequently used to study locomotory activity without visual stimuli 56–58.

Revised version: In recent years, mainly low-throughput, single-fly visual assays like the classical two-stripe Buridan assay (Bülthoff et al., 1982; Götz, 1980), arenas for tethered walking (Seelig et al., 2010) or virtual flight (Mathejczyk & Wernet, 2020) have been used to study behavioral individuality in vison-based tasks (Linneweber et al., 2020; Mathejczyk & Wernet, 2019). In Buridan’s paradigm, flies walk between two inaccessible visual targets (two vertical high-contrast stripes). The stripes are presented in a homogenously illuminated surrounding with no other visual or non-visual cues. Apart from work on individuality involving Dorsal Cluster Neurons (DCNs) (Hassan et al., 2000), also known as Lobula Columnar Neuron 14 (LC14) (Otsuna & Ito, 2006), Buridan’s paradigm has been used to show the importance of the central complex for normal walking speed (Strauss et al., 1992) and for studying photoreceptor function or development (Kiral et al., 2020; Strauss et al., 2001), and visual memory (Neuser et al., 2008). Other circular arenas have been frequently used to study locomotory activity without visual stimuli (Branson et al., 2009; Robie et al., 2017; Valente et al., 2007).

Earlier version: In contrast to these low throughput assays, high throughput assays for studying the behavioral responses of large amounts of animals have been used primarily for genetic screens. Examples include apparatuses to study phototaxis 59, geotaxis 60, olfactory learning 61,62, and circadian rhythms 63. The Ethoscope 64 and MARGO 65 are recent additions to this repertoire that offer several additional high-throughput assays. Except for the Drosophila activity monitors (DAM, Trikinetiks), Ethoscope, and MARGO, these group assays cannot distinguish individual behavioral differences. Only the Ethoscope and MARGO provide some visual modules. Hence, progress towards understanding individuality in visually guided behaviors in a high-throughput fashion remains meager, with few exceptions 9,10,12.

Revised version: In contrast to these low throughput assays, high throughput assays for studying the behavioral responses of large amounts of animals have been used primarily for genetic screens. Examples include apparatuses to study phototaxis (Hirsch & Boudreau, 1958), geotaxis (Hirsch, 1959), olfactory learning (Jiang et al., 2016; Quinn et al., 1974), and circadian rhythms (Konopka & Benzer, 1971). The Ethoscope (Geissmann et al., 2017) and MARGO (Werkhoven et al., 2019) are recent additions to this repertoire that offer several additional high-throughput assays. Except for the Drosophila activity monitors (DAM, Trikinetiks), Ethoscope, and MARGO, these group assays cannot distinguish individual behavioral differences. Only the Ethoscope and MARGO provide some visual modules. Hence, progress towards understanding individuality in visually guided behaviors in a high-throughput fashion remains meager, with few exceptions (Ayroles et al., 2015; Buchanan et al., 2015; Kain et al., 2012).

Earlier version: To investigate the influence of the environment on the stability of visually guided individuality and, to systematically quantify individuality across many behavioral traits and their dependency on the environment in walking flies, we built a new multi-parameter assay called Indytrax. Indytrax is a compact and affordable Mathworks MATLAB-based modular assay and tracker for high-throughput individual ethological quantification. The assay includes swappable modules for multiple high-throughput Buridan paradigms, a Y-maze choice assay for studying the decision-making and locomotor handedness 10, and a simple setup to quantify locomotor activity in large groups of flies 57. We used this tracker to determine the relative contribution of various environmental features (temperature, visual cues, arena shape) to the stability of individual behavior. Our results show that identically defined visual stimuli (stripe color, width, and length) do not necessarily evoke the same behavior in an individual when presented in different environments. Conversely, we find that specific individual behavioral traits (attention, exploration, and anxiety, categorized via multiple behavioral parameters) remain stable even in different environments. Finally, using an LED-based virtual flight arena, we show that attention towards a stimulus remains stable even across different modalities of locomotion (walking versus flying) when using the same Buridan-type visual stimulus. Hence, we find stable behavioral traits across situations in our fly model, in analogy to what has been demonstrated for humans 30,33. Altogether, our results provide the foundation for understanding the complex interaction between individuality and the environment, including the question of why individual behavior is more unpredictable in variable naturalistic conditions, much like the stability across situations in human subjects 19.

Revised version: To investigate the influence of the environment on the consistency of visually guided individuality and, to systematically quantify individuality across many behavioral traits and their dependency on the environment in walking flies, we built a new multi-parameter assay called Indytrax. Indytrax is a compact and affordable Mathworks MATLAB-based modular assay and tracker for high-throughput individual ethological quantification. The assay includes swappable modules for multiple high-throughput Buridan paradigms, a Y-maze choice assay for studying the decision-making and locomotor handedness (Buchanan et al., 2015), and a simple setup to quantify locomotor activity in large groups of flies (Branson et al., 2009). We used this tracker to determine the relative contribution of various environmental features (temperature, visual cues, arena shape) to the consistency of individual behavior. Our results show that identically defined visual stimuli (stripe color, width, and length) do not necessarily evoke the same behavior in an individual when presented in different environments. Conversely, we find that specific individual behavioral traits (attention, exploration, and anxiety, categorized via multiple behavioral parameters) remain stable even in different environments. Finally, using an LED-based virtual flight arena, we show that attention towards a stimulus remains stable even across different modalities of locomotion (walking versus flying) when using the same Buridan-type visual stimulus. Hence, we find stable behavioral traits across situations in our fly model, in analogy to what has been demonstrated for humans (Costa & McCrae, 1988; Fleeson, 2001). Altogether, our results provide the foundation for understanding the complex interaction between individuality and the environment, including the question of why individual behavior is more unpredictable in variable naturalistic conditions, much like the consistency across situations in human subjects (Mischel, 1968).

Materials and Methods

Animals

Earlier version: All experiments were performed using one week old wild-type Drosophila melanogaster (CantonS (Source: Bassem Hassan) and TopBanana 66) reared on a standard Drosophila yeast cornmeal medium (7.5g Agar, 64g cornmeal, 160g yeast, 85.5ml sugar cane syrup in 1L of water) at 25°C and 50% relative humidity under a 12/12 light/dark rhythm. Individual flies were placed into the behavioral arenas using a fine brush and were allowed to accommodate for about 10 minutes before any experiment. The behavioral arenas and flies were cooled on ice for Y-maze experiments before gently loading single flies into the chambers using a fine brush. For experiments with clipped wings, the wings were clipped close to the hinge under CO2 anesthesia, and flies were tested after being allowed to recover within individual vials for at least two days.

Revised version: All experiments were performed using one week old wild-type Drosophila melanogaster (CantonS (Source: Bassem Hassan) and TopBanana (Warren et al., 2018)) reared on a standard Drosophila yeast cornmeal medium (7.5g Agar, 64g cornmeal, 160g yeast, 85.5ml sugar cane syrup in 1L of water) at 25°C and 50% relative humidity under a 12/12 light/dark rhythm. Individual flies were placed into the behavioral arenas using a fine brush and were allowed to accommodate for about 10 minutes before any experiment. The behavioral arenas and flies were cooled on ice for Y-maze experiments before gently loading single flies into the chambers using a fine brush. For experiments with clipped wings, the wings were clipped close to the hinge under CO2 anesthesia, and flies were tested after being allowed to recover within individual vials for at least two days.

Setups

Setup 1: Standard Buridan Arena

Assay

Earlier version: Two Buridan setups were used: type A uses fluorescent tubes for background illumination (as described in 13,67). Type B utilizes a panoramical LED matrix (RGB, 384×128 pixels, pixel size 0.9°, 120Hz frame rate, 5kHz refresh rate, < 1ms Pixel update latency), which is easily addressable via HDMI for stimulus presentation (see Fig. 1a). For a detailed description of the assays and tracking procedures, see supplemental Fig. S1. Flies were walking on a circular platform (diameter 12cm) between two opposing black vertical stripes (angular extent 12° horizontal and 90° vertical measured from the center of the platform). A trench of water surrounded the platform to prevent the flies from escaping (type A) or a clear plastic hemispherical dome coated with Sigmacote (Sigma-Aldrich, SL2-25ML) from the inside (type B). Single-fly walking behavior was recorded at 25°C and 50% relative humidity for 15 minutes (450×450, 8-bit greyscale, 20 Hz, AVI, Mjpeg compression, 90% quality). Before changing an experimental condition, flies were placed back into single vials for recovery.

Revised version: Two Buridan setups were used: type A uses fluorescent tubes for background illumination (as described in (Colomb et al., 2012; Linneweber et al., 2020)). Type B utilizes a panoramical LED matrix (RGB, 384x128 pixels, pixel size 0.9°, 120Hz frame rate, 5kHz refresh rate, < 1ms Pixel update latency), which is easily addressable via HDMI for stimulus presentation (see Fig. 1a). For a detailed description of the assays and tracking procedures, see supplemental Fig. S1. Flies were walking on a circular platform (diameter 12cm) between two opposing black vertical stripes (angular extent 12° horizontal and 90° vertical measured from the center of the platform). A trench of water surrounded the platform to prevent the flies from escaping (type A) or a clear plastic hemispherical dome coated with Sigmacote (Sigma-Aldrich, SL2-25ML) from the inside (type B). Single-fly walking behavior was recorded at 25°C and 50% relative humidity for 15 minutes (450x450, 8-bit greyscale, 20 Hz, AVI, Mjpeg compression, 90% quality). Before changing an experimental condition, flies were placed back into single vials for recovery.

Setups for studying individuality walking flies.

a Overview of a classical (left, fluorescent tubes) and LED (middle) Buridan assay (behavioral platform diameter = 120mm). A MATLAB script provides single fly tracking and analysis of behavioral key components (right). b-d persistence of individual walking patterns across time (b), different stripe numbers (c), and changes in background illumination (d).

Automated tracking procedure

Earlier version: Data was analyzed either in real-time using previously described code (type A, see 67) or offline using a custom-written MATLAB code (type B, available for download from github.com/LinneweberLab/individuality-assays, see supplemental Fig. S1c), respectively. Thresholded video frames and overlaying trajectories are displayed continuously in a separate window to allow the experimenter to assess tracking quality. The MATLAB code works as follows: the circular region of interest (ROI) is automatically detected based on adaptive thresholding and subsequent filtering of the thresholded ROI based on area. Once the ROI is detected, a background image is generated by calculating the maximum intensity between a recorded video’s first and last frames. All flies that moved between those frames will not appear in this background-only image. During tracking, the background image is subtracted from each consecutive frame to simplify the thresholding of the fly. For each frame, the centroid of these thresholded pixels is stored. Should more than one object be detected within a ROI (e.g., through pixel noise), the x-/y-coordinates of the object closest to the fly’s position in the previous frame are stored. If a fly is lost, the tracker keeps the last known position of the fly until it is detected again.

Revised version: Data was analyzed either in real-time using previously described code (type A, see (Colomb et al., 2012)) or offline using a custom-written MATLAB code (type B, available for download from github.com/LinneweberLab/Mathejczyk_2024_eLife_Individuality, see supplemental Fig. S1c), respectively. Thresholded video frames and overlaying trajectories are displayed continuously in a separate window to allow the experimenter to assess tracking quality. The MATLAB code works as follows: the circular region of interest (ROI) is automatically detected based on adaptive thresholding and subsequent filtering of the thresholded ROI based on area. Once the ROI is detected, a background image is generated by calculating the maximum intensity between a recorded video’s first and last frames. All flies that moved between those frames will not appear in this background-only image. During tracking, the background image is subtracted from each consecutive frame to simplify the thresholding of the fly. For each frame, the centroid of these thresholded pixels is stored. Should more than one object be detected within a ROI (e.g., through pixel noise), the x-/y- coordinates of the object closest to the fly’s position in the previous frame are stored. If a fly is lost, the tracker keeps the last known position of the fly until it is detected again.

Data output

Earlier version: Once tracking is completed, the raw tracking results are converted to real-world units (mm) and saved into a .mat-file. Before further analysis, minimal movements are filtered out by only updating the fly’s position if a certain walking threshold (0.8mm/frame) is detected. For each fly, a figure can be saved, allowing for a quick overview of the walking trajectory, and a variety of behavioral parameters are also saved into a separate .txt-file to simplify further data analysis. For a detailed explanation of calculated parameters, see supplemental table 1. For an in-depth analysis of Buridan data, sophisticated data analysis routines describing many behavioral parameters were previously developed 67. To simplify data integration and compatibility, the raw trajectories can also be saved to a pair of .TXT and .DAT files which can later be read and analyzed by previously available software 67.

Revised version: Once tracking is completed, the raw tracking results are converted to real-world units (mm) and saved into a .mat-file. Before further analysis, minimal movements are filtered out by only updating the fly’s position if a certain walking threshold (0.8mm/frame) is detected. For each fly, a figure can be saved, allowing for a quick overview of the walking trajectory, and a variety of behavioral parameters are also saved into a separate .txt-file to simplify further data analysis. For a detailed explanation of calculated parameters, see supplemental table 1. For an in-depth analysis of Buridan data, sophisticated data analysis routines describing many behavioral parameters were previously developed (Colomb et al., 2012). To simplify data integration and compatibility, the raw trajectories can also be saved to a pair of .TXT and .DAT files which can later be read and analyzed by previously available software (Colomb et al., 2012).

Setup 2: IndyTrax multi-arena platform

Assay

Earlier version: To allow for high-throughput behavioral quantification under various tightly controlled environmental contexts, we designed the novel assay system IndyTrax. In short: single flies are transferred to an array of small behavioral arenas. They are filmed from below against near-infrared (850nm + Lee filters 87C IR-pass filter on camera objective) LED illumination (Fig. 3). Compared to published Buridan assays 67, this allows tracking flies in the absence of visible light, but also has the advantage that flies cannot see the camera lens which might otherwise appear as a high-contrast object from within the arena. The assay was designed to allow for the tracking of flies in various possible arena shapes and numbers. We created three different 3D-printable arena types (supplemental Fig. S4e-g), which can be downloaded from github.com/LinneweberLab/individuality-assays. For individuality analysis, we used the multi-Buridan and Y-maze arenas. The setup was built from commercially available and 3D-printed parts. For a part list, see the supplemental table 2. All 3D-printed parts and behavioral assays were printed on a Prusa i3 mk3 printer using black or white PLA filament (Verbatim 55318) and can be downloaded from github.com/LinneweberLab/individuality-assays. Similarly, the MATLAB code tracking and data analysis can be downloaded from there. For detailed instructions on how to install and run the code, see supplemental information. With a computer running MATLAB already available, we estimate the building cost for a complete, temperature- and humidity-controlled enclosure, including all behavioral assays, at about 3500€.

Revised version: To allow for high-throughput behavioral quantification under various tightly controlled environmental contexts, we designed the novel assay system IndyTrax. In short: single flies are transferred to an array of small behavioral arenas. They are filmed from below against near-infrared (850nm + Lee filters 87C IR-pass filter on camera objective) LED illumination (Fig. 3). Compared to published Buridan assays (Colomb et al., 2012), this allows tracking flies in the absence of visible light, but also has the advantage that flies cannot see the camera lens which might otherwise appear as a high-contrast object from within the arena. The assay was designed to allow for the tracking of flies in various possible arena shapes and numbers. We created three different 3D-printable arena types (supplemental Fig. S4e-g), which can be downloaded from github.com/LinneweberLab/Mathejczyk_2024_eLife_Individuality. For individuality analysis, we used the multi-Buridan and Y-maze arenas. The setup was built from commercially available and 3D-printed parts. For a part list, see the supplemental table 2. All 3D-printed parts and behavioral assays were printed on a Prusa i3 mk3 printer using black or white PLA filament (Verbatim 55318) and can be downloaded from github.com/LinneweberLab/Mathejczyk_2024_eLife_Individuality. Similarly, the MATLAB code tracking and data analysis can be downloaded from there. For detailed instructions on how to install and run the code, see supplemental information. With a computer running MATLAB already available, we estimate the building cost for a complete, temperature-and humidity-controlled enclosure, including all behavioral assays, at about 3500€.

Temperature- and humidity-controlled enclosure

Behavioral arenas are located within a black climate-controlled metal enclosure. Temperature and humidity within the enclosure can be regulated using two 3D-printed humidifiers and a heating plate (see supplemental Fig. S4a). Ambient temperature and humidity can be set at any desired value using this configuration and kept stable over a long time (see supplemental Fig. S4c). Since small enclosed arenas are prone to heat build-up, we placed all temperature- and humidity-sensors within the center arena (arena 13) for Buridan experiments and close to the flies in Y-maze experiments. For Buridan experiments, this limits the number of simultaneously trackable arenas to 24 but allows more accurate climate control within the chambers.

Small-scale Buridan-stimuli arrays

Earlier version: We modified the basic design of previously published behavioral arenas optimized for Buridan object orientation behavior 67 by implementing changes in four major areas to allow for high throughput tracking: Firstly, we reduced the arena diameter to allow for filming an array (5×5) of optically isolated Buridan arenas with only a single camera (Fig. 3). Secondly, flies are walking on circular translucent acrylic plates (49mm diameter) which allows for filming flies from below against infrared back-illumination, hence avoiding problems with dead angles that would arise when filming cylindrical arrays from above. Thirdly, during experiments, flies are contained on their respective acrylic platform by clear hemispherical plastic domes (Fig. 3b). This removes the necessity of cutting the flies’ wings before experiments and saves space by making the water-filled moats of the original design superfluous. Plastic domes were coated with Sigmacote (Sigma-Aldrich, SL2-25ML) from the inside to prevent the flies from sitting on the dome surface. Fourthly, to provide a scalable, homogeneously illuminated white background for optical stimulus presentation, the fluorescent tubes of the original design were replaced with electroluminescent (el-)foil wrapped around translucent acrylic cylinders (Fig. 3b). Such el-foils can be cut to size and have low heat emission. They emit light evenly over the whole surface, minimizing unwanted intensity cues within the arenas. El foils are driven at 2 kHz, thus providing virtually flicker-free white illumination to the flies. For creating the Buridan object orientation stimulus, thin black stripes (12° wide as measured from the arena center) were 3D printed and glued vertically to the inside of the arena walls 180° apart, thereby creating high-contrast objects against a uniform white background (Fig. 3b). Individual arenas were fitted into a 5×5 array. Every other arena within the array was rotated by 90° counter-clockwise to increase the robustness of behavioral readouts by canceling out possible directional artifacts created within the enclosure.

Revised version: We modified the basic design of previously published behavioral arenas optimized for Buridan object orientation behavior (Colomb et al., 2012) by implementing changes in four major areas to allow for high throughput tracking: Firstly, we reduced the arena diameter to allow for filming an array (5x5) of optically isolated Buridan arenas with only a single camera (Fig. 3). Secondly, flies are walking on circular translucent acrylic plates (49mm diameter) which allows for filming flies from below against infrared back-illumination, hence avoiding problems with dead angles that would arise when filming cylindrical arrays from above. Thirdly, during experiments, flies are contained on their respective acrylic platform by clear hemispherical plastic domes (Fig. 3b). This removes the necessity of cutting the flies’ wings before experiments and saves space by making the water-filled moats of the original design superfluous. Plastic domes were coated with Sigmacote (Sigma-Aldrich, SL2-25ML) from the inside to prevent the flies from sitting on the dome surface. Fourthly, to provide a scalable, homogeneously illuminated white background for optical stimulus presentation, the fluorescent tubes of the original design were replaced with electroluminescent (el-)foil wrapped around translucent acrylic cylinders (Fig. 3b). Such el-foils can be cut to size and have low heat emission. They emit light evenly over the whole surface, minimizing unwanted intensity cues within the arenas. El foils are driven at 2 kHz, thus providing virtually flicker-free white illumination to the flies. For creating the Buridan object orientation stimulus, thin black stripes (12° wide as measured from the arena center) were 3D printed and glued vertically to the inside of the arena walls 180° apart, thereby creating high-contrast objects against a uniform white background (Fig. 3b). Individual arenas were fitted into a 5x5 array. Every other arena within the array was rotated by 90° counter-clockwise to increase the robustness of behavioral readouts by canceling out possible directional artifacts created within the enclosure.

Earlier version: Walking behavior was recorded at 50% relative humidity for 15 minutes (1200×1200 pixel, 8-bit greyscale, 15 Hz, AVI, Mjpeg compression, 90% quality) with either the el-foil-illumination on, in darkness, at 23°C or 32°C, respectively. After being tested in the Buridan assay, the flies were put back into single vials overnight to allow for recovery before testing them in the Y-maze assay the next day around the same time.

Revised version: Walking behavior was recorded at 50% relative humidity for 15 minutes (1200x1200 pixel, 8-bit greyscale, 15 Hz, AVI, Mjpeg compression, 90% quality) with either the el-foil-illumination on, in darkness, at 23°C or 32°C, respectively. After being tested in the Buridan assay, the flies were put back into single vials overnight to allow for recovery before testing them in the Y-maze assay the next day around the same time.

Y-maze array

Earlier version: We constructed a modular 3D-printable Y-maze array (144 arenas, 300×300×2mm) based on a previously published design 65. The Y-maze plates were placed between two 320×320×3mm borosilicate glass plates. The bottom side of the upper glass plate was coated three times using Sigmacote (Sigma-Aldrich, SL2-25ML) to prevent the flies from walking on the upper plate (Fig. 3c). Directly below the arenas and above the upper glass plate, we placed a layer of white paper (Evercopy premium 80g, Ref. 1902C) for light diffusion and to restrict the view of the flies outside of the arenas. For individuality experiments, we covered the outer arenas with black tape and tracked flies only within 100 arenas located in the center. This made keeping track of individual fly identities easier when switching between the different behavioral assays. Videos of walking flies were recorded at 50% relative humidity for 30 minutes (1200×1200 pixel, 8-bit greyscale, 15 Hz, AVI, Mjpeg compression, 90% quality) under white light illumination, in darkness, at 23°C or 32°C, respectively. For white light illumination, we placed a 30×30cm LED panel (Tween Light 30 x 30 cm LED panel, 4000k, 16W) on top of the arenas and temporarily removed the infrared-pass filter from the camera objective. Y-maze experiments were all done consecutively on the same day (light on 23°C, light off 23°C, light on 32°C, light off 32°C).

Revised version: We constructed a modular 3D-printable Y-maze array (144 arenas, 300x300x2mm) based on a previously published design (Werkhoven et al., 2019). The Y-maze plates were placed between two 320x320x3mm borosilicate glass plates. The bottom side of the upper glass plate was coated three times using Sigmacote (Sigma-Aldrich, SL2-25ML) to prevent the flies from walking on the upper plate (Fig. 3c). Directly below the arenas and above the upper glass plate, we placed a layer of white paper (Evercopy premium 80g, Ref. 1902C) for light diffusion and to restrict the view of the flies outside of the arenas. For individuality experiments, we covered the outer arenas with black tape and tracked flies only within 100 arenas located in the center. This made keeping track of individual fly identities easier when switching between the different behavioral assays. Videos of walking flies were recorded at 50% relative humidity for 30 minutes (1200x1200 pixel, 8-bit greyscale, 15 Hz, AVI, Mjpeg compression, 90% quality) under white light illumination, in darkness, at 23°C or 32°C, respectively. For white light illumination, we placed a 30x30cm LED panel (Tween Light 30 x 30 cm LED panel, 4000k, 16W) on top of the arenas and temporarily removed the infrared-pass filter from the camera objective. Y-maze experiments were all done consecutively on the same day (light on 23°C, light off 23°C, light on 32°C, light off 32°C).

Automated tracking procedure

We developed a fully-automated MATLAB-based tracking code that allows users to adjust parameters like arena number and size and generates output data for all videos within a user-selectable folder. The code assumes only one fly per region of interest. This saves processing time and increases tracking robustness by not having to match multiple fly identities between frames. Before running the code, the experimenter must create a camera calibration file using the MATLAB camera calibration toolbox (see supplemental information). This allows for more accurate tracking results by correcting lens distortions (supplemental Fig. S4d). If the camera stays immobile, this procedure must be done only once, usually only a few minutes. As long as the arenas filmed are arranged in rows and columns, the code can handle virtually any shape or number of arenas, making it ideal for tracking flies across different behavioral assays (supplemental Fig. S4e-g). Regions of interest (ROIs) are detected based on adaptive thresholding and subsequent filtering of thresholded ROI-objects based on the area (supplemental Fig. S4e-g, middle). Once the ROIs are detected and sorted, a background-subtracted image is generated, allowing the user to determine a detection threshold, above which all pixels will be interpreted as foreground objects (supplemental Fig. S4e, third from left). During tracking, background subtraction and object thresholding is performed for each frame, and for each ROI the centroid of the thresholded fly is extracted and saved. Thresholded pixels and trajectories are displayed continuously in a separate window to allow the experimenter to assess tracking quality. If more than one object is detected within a ROI (e.g. through pixel noise), the x-/y-coordinates of the object closest to the fly’s position in the previous frame is stored. If a fly is lost, the tracker keeps the last known position of the fly until it is detected again. The average tracking speed for 24 ROIs/flies was about 20hz with MATLAB running on an Intel i5 3Ghz CPU. For tracking 400 ROIs/flies average tracking speed was about 5hz.

Data output

Data output is the same as described for our standard Buridan arenas.

Setup 3: Virtual flight simulator

Assay

Earlier version: We designed a novel, easy-to-assemble, and easy-to-use virtual flight simulator using affordable 3D-printable and off-the-shelf parts (Fig. 5a-c) to quantify behavioral responses during flight. For a detailed description of the setup and custom code, see supplemental Fig. S9. In short, using a tethering station (for description, see 48), flies were cooled down for immobilization and glued to a steel pin (0.1mm diameter, 10mm length, www.entosphinx.cz) using UV-cured glue (Bondic). Once tethered, single flies were put between two neodymium magnets (magneto-tether; upper magnet: diameter 5mm, lower magnet: ring, outer diameter 10mm, inner diameter 5mm). Once the appropriate distance between the magnet is found (about 20mm), the magnetic field allows the flies to rotate freely around their yaw-axis while keeping the steel pin in place (Fig. 5b). The tip of the steel pin is inserted into a V-shaped sapphire bearing (1mm diameter) to reduce friction. For re-initiation of flight, air puffs can be delivered from below via a membrane pump. A panoramic LED matrix surrounding the magneto-tether was used for Buridan stimulus presentation (RGB, 256×128 pixels, 120Hz frame rate, 5kHz refresh rate, pixel size 1.4°, angular extent 360° horizontal x 100° vertical, stripe width: 12°). The LED matrix controller is connected to a computer via HDMI, allowing an easy and flexible stimulus creation and presentation (Novastar mctrl 660 pro, 120Hz, < 1ms Pixel update latency). Flies were filmed from below under near-infrared illumination (Fig. 5c), and each fly’s heading was tracked in real-time (90 Hz) using a custom-written MATLAB code (available for download from github.com/LinneweberLab/individuality-assays). Before starting the experiment, each fly was tested for 120s with an optomotor stimulus (rotating stripe pattern, 120°/s, 40x CW and 40x CCW, respectively) to assess the flies’ ability to rotate smoothly (Fig. 5d, left). After this, flies were presented with the Buridan stimulus with varying contrast (100%, 50%, 100% inverted; stripe width 30°), and their heading was tracked for 3×4 min. Flies were then carefully separated from the steel pin and transferred into single vials for overnight recovery before being re-tethered and tested the following day around the same time. For flight vs. walking experiments, flies were tested for 10 min while flying under a 100% contrast Buridan stimulus (stripe width 12° from a fly’s perspective) before being gently separated from the steel pin and being transferred into the LED Buridan assay (Type B), where they were tested for 15min under a 100% contrast Buridan stimulus (stripe width 12° from a fly’s perspective).

Revised version: We designed a novel, easy-to-assemble, and easy-to-use virtual flight simulator using affordable 3D-printable and off-the-shelf parts (Fig. 5a-c) to quantify behavioral responses during flight. For a detailed description of the setup and custom code, see supplemental Fig. S9. In short, using a tethering station (for description, see (Mathejczyk & Wernet, 2020)), flies were cooled down for immobilization and glued to a steel pin (0.1mm diameter, 10mm length, www.entosphinx.cz) using UV-cured glue (Bondic). Once tethered, single flies were put between two neodymium magnets (magneto-tether; upper magnet: diameter 5mm, lower magnet: ring, outer diameter 10mm, inner diameter 5mm). Once the appropriate distance between the magnet is found (about 20mm), the magnetic field allows the flies to rotate freely around their yaw-axis while keeping the steel pin in place (Fig. 5b). The tip of the steel pin is inserted into a V-shaped sapphire bearing (1mm diameter) to reduce friction. For re-initiation of flight, air puffs can be delivered from below via a membrane pump. A panoramic LED matrix surrounding the magneto-tether was used for Buridan stimulus presentation (RGB, 256x128 pixels, 120Hz frame rate, 5kHz refresh rate, pixel size 1.4°, angular extent 360° horizontal x 100° vertical, stripe width: 12°). The LED matrix controller is connected to a computer via HDMI, allowing an easy and flexible stimulus creation and presentation (Novastar mctrl 660 pro, 120Hz, < 1ms Pixel update latency). Flies were filmed from below under near-infrared illumination (Fig. 5c), and each fly’s heading was tracked in real-time (90 Hz) using a custom-written MATLAB code (available for download from github.com/LinneweberLab/Mathejczyk_2024_eLife_Individuality). Before starting the experiment, each fly was tested for 120s with an optomotor stimulus (rotating stripe pattern, 120°/s, 40x CW and 40x CCW, respectively) to assess the flies’ ability to rotate smoothly (Fig. 5d, left). After this, flies were presented with the Buridan stimulus with varying contrast (100%, 50%, 100% inverted; stripe width 30°), and their heading was tracked for 3x4 min. Flies were then carefully separated from the steel pin and transferred into single vials for overnight recovery before being re-tethered and tested the following day around the same time. For flight vs. walking experiments, flies were tested for 10 min while flying under a 100% contrast Buridan stimulus (stripe width 12° from a fly’s perspective) before being gently separated from the steel pin and being transferred into the LED Buridan assay (Type B), where they were tested for 15min under a 100% contrast Buridan stimulus (stripe width 12° from a fly’s perspective).

Automated tracking procedure

Earlier version: The code for tracking a fly’s heading in 360° can be downloaded from github.com/LinneweberLab/individuality-assays. Flight tracking is done in real-time at 90fps (see supplemental Fig. S9c). A timestamp is created for each incoming camera frame, and two different thresholded images are computed with different Gaussian filter settings. A strong filter setting blurs out the legs of the fly and allows for thresholding only the fly’s body and calculating a heading angle between 0° and 180° through ellipse fitting. By subtracting the thresholded body pixels from a weaker filtered version of the same frame, the limbs and wings (visible when the fly stops flying) can also be tracked separately if the experimenter wants. To calculate the heading between 0° and 360°, the centroid of thresholded body pixels is set to (0,0) within a fictive coordinate system (see supplemental Fig. S9c), and the body axis is aligned with the x-axis. The head region is assigned based on the sum of pixels within the 4 axis quadrants (smallest area of either I+IV or II+III, respectively). Once the heading is calculated, it is stored together with the original video frame for additional offline data analysis, if desired. When a fly stops flying (wings get detected in the video), a command is sent to a connected Arduino, which controls a relay to switch on a membrane pump and provide an air puff to reinitiate flight. During tracking, the thresholded video is displayed together with the computed heading to allow the experimenter to assess tracking quality.

Revised version: The code for tracking a fly’s heading in 360° can be downloaded from github.com/LinneweberLab/Mathejczyk_2024_eLife_Individuality. Flight tracking is done in real-time at 90fps (see supplemental Fig. S9c). A timestamp is created for each incoming camera frame, and two different thresholded images are computed with different Gaussian filter settings. A strong filter setting blurs out the legs of the fly and allows for thresholding only the fly’s body and calculating a heading angle between 0° and 180° through ellipse fitting. By subtracting the thresholded body pixels from a weaker filtered version of the same frame, the limbs and wings (visible when the fly stops flying) can also be tracked separately if the experimenter wants. To calculate the heading between 0° and 360°, the centroid of thresholded body pixels is set to (0,0) within a fictive coordinate system (see supplemental Fig. S9c), and the body axis is aligned with the x-axis. The head region is assigned based on the sum of pixels within the 4 axis quadrants (smallest area of either I+IV or II+III, respectively). Once the heading is calculated, it is stored together with the original video frame for additional offline data analysis, if desired. When a fly stops flying (wings get detected in the video), a command is sent to a connected Arduino, which controls a relay to switch on a membrane pump and provide an air puff to reinitiate flight. During tracking, the thresholded video is displayed together with the computed heading to allow the experimenter to assess tracking quality.

Data output

After tracking, raw heading data, a graphical overview image, and a table containing computed behavioral parameters are saved automatically for each fly which can be used for further data analysis. For a detailed explanation of all calculated parameters, see supplemental table 1.

Added in the revised version: Data and statistical analysis for all experiments

Removed in the revised version: Data analysis

Earlier version: All data analysis was performed with MATLAB. For angular calculations, we used algorithm presented in 68. After a test for normality (Shapiro-Wilk-Test), behavioral data between groups was compared using the paired T-test. For correlating individual data, we calculated Pearson correlation coefficients. For walking data, an arena’s center region (edge-corrected area) was defined as the circular area around the center with a radius of 80% of the arena radius. The maximum color value (red) for generating heatmaps was set to the 95%-quantile of the count distribution to increase data visibility.

Revised version: All data analysis was performed with MATLAB. For angular calculations, we used algorithm presented in (Berens, 2009). After a test for normality (Shapiro-Wilk-Test), behavioral data between groups was compared using the paired T-test. For correlating individual data, we calculated Pearson correlation coefficients. For walking data, an arena’s center region (edge-corrected area) was defined as the circular area around the center with a radius of 80% of the arena radius. The maximum color value (red) for generating heatmaps was set to the 95%-quantile of the count distribution to increase data visibility. Flies that did not move during the experiment were excluded from the analysis.

Added in the revised version: Generalized Linear Model (GLM)

Added in the revised version: To summarize and generalize the findings from all experiments in this study (using individual rank differences upon changing one or more environmental parameters), a Generalized Linear Model (GLM) was trained using the environmental parameters as predictors (0 when environmental parameter was not changed and 1 if it was changed for any given experimental group) and the resulting individual rank differences as the response variable. We employed ridge regression to mitigate collinearity among predictors and to control for overfitting. A range of regularization parameters (λ) was tested using 5-fold cross-validation to identify the optimal value that minimized the mean squared error (MSE). Z-scoring was applied to the predictor variables prior to model fitting for data normalization. To estimate coefficient variability and assess significance, bootstrapping was applied. For each of the 1,000 bootstrap iterations, predictors and response data were resampled with replacement, and the ridge regression model was refit using the optimal λ. Coefficients and intercept values from each bootstrap sample were stored, and their distributions were used to calculate the mean coefficients, standard errors (SEs) and p-values for coefficients based on t-statistics derived from bootstrapped SEs.

Results

Earlier version: The goal of our study was to understand to what extent animal individuality is influenced by situational changes in the environment, i.e., how much of an animal’s individuality remains after one or more environmental features change. Based on our previous work on animal individuality 13, we started our analysis in the multiparametric two-stripe Buridan’s paradigm (Fig. 1a, supplemental Fig. S1 for technical details) 13,67. As expected, we measured interindividually variable, yet remarkably stable behavioral responses, both on the individual (Fig. 1b, Fig. 2a, supplemental Fig. S3a) and mean group level (supplemental Fig. S2a), when tested in the exact same situation with all environmental features kept the same. This was true for the five representative behavioral parameters (out of 20 parameters total, see supplemental Table T1 and supplemental Fig. S8), defined as falling into three behavioral traits: exploration (% of time walked, walking speed), attention (vector strength, angular velocity), and anxiety (centrophobicity).

Revised version: The goal of our study was to understand to what extent animal individuality is influenced by situational changes in the environment, i.e., how much of an animal’s individuality remains after one or more environmental features change. Based on our previous work on animal individuality (Linneweber et al., 2020), we started our analysis in the multiparametric two-stripe Buridan’s paradigm (Fig. 1a, supplemental Fig. S1 for technical details) (Colomb et al., 2012; Linneweber et al., 2020). As expected, we measured interindividually variable, yet remarkably consistent behavioral responses, both on the individual (Fig. 1b, Fig. 2a, supplemental Fig. S3a) and mean group level (average of all individuals in the same situation, supplemental Fig. S2a), when tested in the exact same situation with all environmental features kept the same. This was true for the five representative behavioral parameters (out of 20 parameters total, see supplemental Table T1 and supplemental Fig. S8), defined in three behavioral traits: exploration (% of time walked, walking speed), attention (vector strength, angular velocity), and anxiety (centrophobicity). The exploration trait is characterized by parameters describing activity. Parameters related to attention reflect heightened responses to visual cues, but unlike commonly used metrics such as angle or stripe deviations (Colomb et al., 2012; Linneweber et al., 2020), they can be calculated in the absence of visual cues. Lastly, the parameter centrophobicity, used as an indicator of the trait anxiety, is linked to the open-field test in mice, where the ratio of wall-to-open-field activity is frequently calculated as a measurement of anxiety (Carter & Shieh, 2015).

Individual traits of walking flies persist over time but depend on the visual scenery.

Earlier version: a-c Correlation of key behavioral parameters over time (a), different stripe numbers (b), and changes in background illumination (c). Correlation coefficients are displayed in the upper left corner of each plot. *=p<0.05, **=p<0.01, ***=p<0.001, for detailed fly numbers per experimental condition, see supplemental Fig. S2.

Revised version: a-c Correlation of key behavioral parameters over time (n=60) (a), different stripe numbers (n=57) (b), and changes in background illumination (n=54) (c). Correlation coefficients are displayed in the upper left corner of each plot. *=p<0.05, **=p<0.01, ***=p<0.001, for detailed fly numbers per experimental condition, see supplemental Fig. S2.

Earlier version: Modifying the visual stimulus within the Buridan arena (0, 1, and 2 stripes were tested, Fig. 1c, Fig. 2b, and supplemental Fig. S2b and S3b) significantly affected the mean group behavior across all three behavioral traits (supplemental Fig. S2b). Especially the population means of three parameters (% of time walked, vector strength, and centrophobicity) were strongly affected by the stimulus modification. Surprisingly, despite this change in population means, we observed only a minor effect on individuality (Fig. 1c and Fig. 2b, measured by a high Pearson correlation coefficient and a consistent rank within the distribution). The individuality for the traits exploration and anxiety persisted even upon modulation (variation of stripe number) and after removing visual cues from the arena. Only those parameters describing the fly’s attention towards visual cues showed an unsurprising reduction in behavioral stability upon removal of the visual stimuli (Fig. 2b). Altogether, these results demonstrate that individuality persisted after stimulus modifications within the same arena, even when the modifications resulted in altered population responses.

Revised version: Modifying the visual stimulus within the Buridan arena (0, 1, and 2 stripes were tested, Fig. 1c, Fig. 2b, and supplemental Fig. S2b and S3b) significantly affected the mean group behavior across all three behavioral traits (supplemental Fig. S2b). Especially the means of three parameters (% of time walked, vector strength, and centrophobicity) were strongly affected by the stimulus modification. Surprisingly, despite this change in group means, we observed only a minor effect on individuality when testing the same fly in different situations (Fig. 1c and Fig. 2b, measured by a high Pearson correlation coefficient and a consistent rank within the distribution). The individuality for the traits exploration and anxiety persisted even upon modulation (variation of stripe number) and after removing visual cues from the arena. Only those parameters describing the fly’s attention towards visual cues showed an unsurprising reduction in behavioral consistency upon removal of the visual stimuli (Fig. 2b). Altogether, these results demonstrate that individuality persisted after stimulus modifications within the same arena, even when the modifications resulted in altered group responses.

Earlier version: To further characterize the surprising relationship between individuality and population means upon situational changes, we modified the homogeneity of background illumination in the Buridan arena using two different diffuser materials (paper and acryl-based covers of fluorescent ring lights) and compared these conditions to a newly designed Buridan LED arena displaying the same visual stimulus (stripes of the same color, width, and height) (Fig. 1d, Fig. 2c and supplemental Fig. S2c and S3c). We measured no significant differences between the population means of the two diffuser conditions and only minor group-based changes compared to the LED Buridan (supplemental Fig. S2c). We then tested the correlation of individual behavior across these conditions, as a measurement of behavioral stability, expecting high correlations. Indeed, the different diffusers within the same arenas only had minor effects on the behavior of individual flies (Fig. 1d). Conversely, the individual response changed dramatically when the same fly was tested again in the LED arena (Fig. 1d). A fly’s attention and anxiety parameters quantified within the original arena did not allow for any prediction for the LED arena (Fig. 2c). Significant correlations between the LED arena and the original arenas were only measured for exploration parameters (Fig. 2c). This unpredictability of individual fly behavior across situations was once again particularly surprising since group-based performances remained similar (supplemental Fig. S2c). Finally, a detailed comparison of males and females revealed only quantitative but no qualitative differences across all parameters between the two sexes (supplemental Fig. S2 and supplemental Fig. S3).

Revised version: To further characterize the surprising relationship between individuality and group means upon situational changes, we modified the homogeneity of background illumination in the Buridan arena using two different diffuser materials (paper and acryl-based covers of fluorescent ring lights) and compared these conditions to a newly designed Buridan LED arena displaying the same visual stimulus (stripes of the same color, width, and height) (Fig. 1d, Fig. 2c and supplemental Fig. S2c and S3c). We measured no significant differences between the group means of the two diffuser conditions and only minor group-based changes compared to the LED Buridan (supplemental Fig. S2c). We then tested the correlation of individual behavior across these conditions, as a measurement of behavioral consistency, expecting high correlations. Indeed, the different diffusers within the same arenas only had minor effects on the behavior of individual flies (Fig. 1d). Conversely, the individual response changed dramatically when the same fly was tested again in the LED arena, albeit presenting the same stimulus (Fig. 1d). A fly’s attention and anxiety parameters quantified within the original arena did not allow for any prediction for the LED arena (Fig. 2c). Significant correlations between the LED arena and the original arenas were only measured for exploration parameters (Fig. 2c). This unpredictability of individual fly behavior across situations was once again particularly surprising since group-based performances remained similar (supplemental Fig. S2c). Finally, a detailed comparison of males and females revealed only quantitative but no qualitative differences across all parameters between the two sexes (supplemental Fig. S2 and supplemental Fig. S3).

Earlier version: Based on the surprising result that individuality 1 of visual orientation behavior is more strongly influenced by the arena type (despite producing very similar population responses) and, only to a lesser degree, dependent on the visual cue (despite strong differences in the population response), we set out to characterize the contribution of environmental situations to individuality even more quantitatively. We therefore built a multiplatform high-throughput assay (Indytrax) to quantitatively study the behavior of the same individual flies in different environmental situations (Fig. 3, technical details supplemental Fig. S4).

Revised version: Based on the surprising result that individuality (Gosling, 2008) of visual orientation behavior is more strongly influenced by the arena type (despite producing very similar group responses) and, only to a lesser degree, dependent on the visual cue (despite strong differences in the group response), we set out to characterize the contribution of environmental situations to individuality even more quantitatively. We therefore built a multiplatform high-throughput assay (Indytrax) to quantitatively study the behavior of the same individual flies in different environmental situations (Fig. 3, technical details supplemental Fig. S4).

A high-throughput assay for studying the influence of temperature, illumination, and arena shape on individuality.

Earlier version: a Setup overview. Flies are filmed from below against a near-infrared backlight and are walking within a 5×5 array of small Buridan chambers (b) or one of up to 144 Y-maze arenas (c). The setup is housed within a temperature- and humidity-controlled enclosure (600 x 450 x 820 mm). The provided MATLAB code corrects lens distortions, automatically detects the arenas, and tracks each fly’s position over time, allowing for generating high-throughput single-fly behavioral data (d).

Revised version: a Setup overview. Flies are filmed from below against a near-infrared backlight and are walking within a 5x5 array of small Buridan chambers (b) or one of up to 144 Y-maze arenas (c). The setup is housed within a temperature- and humidity-controlled enclosure (600 x 450 x 820 mm). The provided MATLAB code corrects lens distortions, automatically detects the arenas, and tracks each fly’s position over time, allowing for generating high-throughput single-fly behavioral data (d). See Materials and Methods:Setup 2: IndyTrax multi-arena platform for details.

Earlier version: Using Indytrax, we compared the behavioral parameters of each fly in eight different environmental situations (Fig. 4, supplemental Fig. S6, S7, S8). The flies were tested in two independent arenas (Buridan and Y-maze) at two different temperatures (23°C and 32°C, supplemental Fig. S5) and under both light and dark conditions. Our analysis showed that temperature significantly affected exploration parameters on the population mean level. As expected, most animals tested were more explorative at higher than at lower temperatures (reviewed in: 69) (supplemental Fig. S6). Despite these population changes, individuals kept their rank within the population independent of the changed baseline activity. For example, fly #1 was the least explorative, fly #2 exploration performance ranked medium, and fly #3 the highest (Fig. 4a). The graphical analysis of exploration parameters (Fig. 4b) confirmed that most individuals maintain their rank in the distribution under different temperatures and light-dark conditions, whereas the individual distribution rank changed between the different arena types. This relationship was less pronounced for attention parameters. The detailed correlative analysis between individuals underscores these observations (Fig. 4c and supplemental Fig. S7 and full analysis supplemental Fig. S8): Amongst the exploration parameters, the percentage of time walked significantly correlated between individuals (Pearson r=0.263-0.398) independently of temperature, illumination, and arena. Walking speed was highly correlated between individuals (Pearson r=0.81, 0.825) under temperature and illumination changes but uncorrelated in different arena types (Pearson r=0.091). Vector strength was the most correlated attention parameter. The individual correlations ranged from Pearson r=0.533 to -0197. Angular velocity significantly correlated for temperature and illumination (Pearson r=0.246, 0.354) but not in different arenas (Pearson r=0.087). The same is true for anxiety measured through centrophobicity (Pearson r=0.694, 0.740, 0.033). In summary, these data revealed that individuality measured as three traits is mostly unaffected by temperature and illumination but strongly affected by the arena type. Our analysis confirmed that even in situations with a large effect on the mean group behavior (like temperature), individual behavioral stability of idiosyncrasies remained high, while conversely, other environmental situations like the arena type had little impact on the mean but a rather massive effect on individual behavioral stability (Fig. S6). Qualitatively similar results leading to the same conclusions were obtained for both sexes (Fig. S6a, S7a) and two different genotypes (Fig. S6b, S7b)

Revised version: Using Indytrax, we compared the behavioral parameters of each fly in eight different environmental situations (Fig. 4, supplemental Fig. S6, S7, S8). The flies were tested in two independent arenas (Buridan and Y-maze) at two different temperatures (23°C and 32°C, supplemental Fig. S5) and under both light and dark conditions. Our analysis showed that temperature significantly affected exploration parameters on the group mean level. As expected, most animals tested were more explorative at higher than at lower temperatures ((reviewed in: Gibert et al., 2016)) (supplemental Fig. S6). Despite these group changes, individuals kept their rank within the group independent of the changed baseline activity. For example, fly #1 was the least explorative, fly #2 exploration performance ranked medium, and fly #3 the highest (Fig. 4a). The graphical analysis of exploration parameters (Fig. 4b) confirmed that most individuals maintain their rank in the distribution under different temperatures and light-dark conditions, whereas the individual distribution rank changed between the different arena types. This relationship was less pronounced for attention parameters. The detailed correlative analysis between individuals underscores these observations (Fig. 4c and supplemental Fig. S7 and full analysis supplemental Fig. S8): Amongst the exploration parameters, the percentage of time walked significantly correlated between individuals (Pearson r=0.263-0.398) independently of temperature, illumination, and arena. Walking speed was highly correlated between individuals (Pearson r=0.81, 0.825) under temperature and illumination changes but uncorrelated in different arena types (Pearson r=0.091). Vector strength was the most correlated attention parameter. The individual correlations ranged from Pearson r=0.533 to -0197. Angular velocity significantly correlated for temperature and illumination (Pearson r=0.246, 0.354) but not in different arenas (Pearson r=0.087). The same is true for anxiety measured through centrophobicity (Pearson r=0.694, 0.740, 0.033). In summary, these data revealed that individuality measured in three traits is mostly unaffected by temperature and illumination but strongly affected by the arena type. Our analysis confirmed that even in situations with a large effect on the mean group behavior (like temperature), individual behavioral consistency of idiosyncrasies remained high, while conversely, other environmental situations like the arena type had little impact on the mean but a massive effect on individual behavioral consistency (Fig. S6). Qualitatively similar results leading to the same conclusions were obtained for both sexes (Fig. S6a, S7a) and two different genotypes (Fig. S6b, S7b)

The influence of temperature, illumination, and arena shape on individuality in walking flies.

Earlier version: A Individuality in walking patterns across different temperatures, illumination, and arena type. Diameter circular arena = 49mm, arm length from center Y-maze = 13.5mm. b changes in the % of time walked for all tested flies depending on environmental conditions. c Changes in the % of time walked grouped by temperature, illumination, and arena. d Correlation of key behavioral parameters across temperature, illumination, and arena type. For detailed fly numbers per experimental condition, see supplemental Fig. S6.

Revised version: a Individuality in walking patterns across different temperatures, illumination, and arena type. Diameter circular arena = 49mm, arm length from center Y-maze = 13.5mm. b changes in the % of time walked and vector strength for all tested flies depending on environmental conditions (n=400). Right: data grouped by temperature, illumination, and arena. c Correlation of key behavioral parameters across temperature, illumination, and arena type (n=400). For detailed fly numbers per experimental condition, see supplemental Fig. S6.

Earlier version: Next, we drastically modified the behavioral situation while keeping the visual stimulus as constant as possible. We, therefore, decided to investigate the influence of a fly’s modality of locomotion (walking versus flying) on individuality. To do this, we designed a new virtual flight arena for tethered flight utilizing the exact same LED panels as were used in the LED Buridan arena (Fig. 5a-c and supplemental S9 for technical details), to maximize the similarity of the visual context. This flight simulator tracked a fly’s heading in real-time in response to dynamic or static panoramic visual stimuli with high resolution (1.4° pixel size). This allowed us to quantitatively compare the visual decisions of both walking and flying flies in response to standard optomotor and Buridan-like stimuli (Fig. 5d-e). We first tested individuality in flying flies over time under the same and a set of modified visual stimuli (Fig. 6, supplemental Fig. S10 and S11). Each fly was tested over two consecutive days under three different visual conditions (dark stripes on a bright background with either 50% or 100% contrast, and an inverted Buridan with bright stripes on a dark background and 100% contrast) (Fig. 6a). Overall, the flies showed some individuality in visual attention across days and contrast. This was also true for the inverted Buridan stimulus, in which the flies oriented themselves toward the dark areas, similar to previous results 70. To allow a quantitative comparison across behavioral traits between walking and flight behavior, we derived an adapted set of parameters for exploration and attention quantification (exploration parameters: # of pauses, Absolute angular velocity; attention parameters: Vector strength, Angular velocity, and Median heading axial, Fig. 6b). The resulting analyses revealed individuality in flying flies (Fig. 6c). Only the inversion of the Buridan stimulus resulted in variable angles towards the black background on different days. However, it is remarkable that on a given day, each fly selected an arbitrary heading that was kept for the entire experimental length (Fig. 6a). Again, the quantitative analysis between males and females did not reveal qualitative differences (Fig. S11).

Revised version: Next, we drastically modified the behavioral situation while keeping the visual stimulus as constant as possible. We, therefore, decided to investigate the influence of a fly’s modality of locomotion (walking versus flying) on individuality. To do this, we designed a new virtual flight arena for tethered flight utilizing the exact same LED panels as were used in the LED Buridan arena (Fig. 5a-c and supplemental S9 for technical details), to maximize the similarity of the visual context. This flight simulator tracked a fly’s heading in real-time in response to dynamic or static panoramic visual stimuli with high resolution (1.4° pixel size). This allowed us to quantitatively compare the visual decisions of both walking and flying flies in response to standard optomotor and Buridan-like stimuli (Fig. 5d-e). We first tested individuality in flying flies over time under the same and a set of modified visual stimuli (Fig. 6, supplemental Fig. S10 and S11). Each fly was tested over two consecutive days under three different visual conditions (dark stripes on a bright background with either 50% or 100% contrast, and an inverted Buridan with bright stripes on a dark background and 100% contrast) (Fig. 6a). Overall, the flies showed some individuality in visual attention across days and contrast. This was also true for the inverted Buridan stimulus, in which the flies oriented themselves toward the dark areas, similar to previous results (Han et al., 2021). To allow a quantitative comparison across behavioral traits between walking and flight behavior, we derived an adapted set of parameters for exploration and attention quantification (exploration parameters: # of pauses, Absolute angular velocity; attention parameters: Vector strength, Angular velocity, and Median heading axial, Fig. 6b). The resulting analyses revealed individuality in flying flies (Fig. 6c). Only the inversion of the Buridan stimulus resulted in variable angles towards the black background on different days. However, it is remarkable that on a given day, each fly selected an arbitrary heading that was kept for the entire experimental length (Fig. 6a). Again, the quantitative analysis between males and females did not reveal qualitative differences (Fig. S11).

A virtual flight simulator for studying individuality in flying flies.

a Setup overview (225 x 225 x 600 mm). Flies are glued to a steel pin and placed in a magneto-tether surrounded by an LED matrix for stimulus presentation. b Close-up photo of a tethered flying fly. c Schematic setup overview. Flies are filmed from below through an IR-pass filter against near-infrared back illumination. A connected pump provides air puffs to reinitiate flight if necessary. We provide a mounting port for an optional IR-laser for heat punishment via a dichroic mirror. d A MATLAB script controls stimulus presentation (optomotor or Buridan stimulus) with an angular extent of 360° azimuthally and 100° vertically and provides simultaneous real-time tracking of flight-heading over 360° (e). Before an experiment, each fly’s ability to rotate was tested using a 120°/s CW or CCW stripe rotation (optomotor stimulus), resulting in measured rotational flight movements close to the rotational velocity of the stripe pattern (e, upper left) and uniform distribution of angles (e, lower left). This was followed by each fly flying under the Buridan stimulus for 10min (d, right) and an analysis of heading choices (e, upper right) and their angular distribution (e, lower right).

Individual traits of flying flies persist over time but depend on the visual scenery.

Earlier version: a Angular histograms show the persistence of individual heading choices over consecutive days and across different contrast ratios. b,c Correlation of key behavioral parameters over time (b) and depending on changes in visual contrast (c). For detailed fly numbers per experimental condition, see supplemental Fig. S10.

Revised version: a Angular histograms show the persistence of individual heading choices over consecutive days and across different contrast ratios. b,c Correlation of key behavioral parameters over time (n=30) (b) and depending on changes in visual contrast (n=30) (c).

Earlier version: After confirming that individuality also existed in flying flies, we tested how different locomotion modalities affected a given animal’s behavioral individuality. Behavioral parameters of both walking and flying flies were quantified in response to the same panoramic LED screen displaying an identical two-stripe Buridan stimulus (Fig. 7a, supplemental Fig. S12 and S13).

Revised version: After confirming that individuality also existed in flying flies, we tested how different locomotion modalities affected a given animal’s behavioral individuality. Behavioral parameters of both walking and flying flies were quantified in response to the same panoramic LED screen displaying an identical two-stripe Buridan stimulus (Fig. 7a, supplemental Fig. S12 and S13). We expected one of three possible outcomes: 1. There is no consistency of individual behavioral traits across flying and walking, 2. Only some behavioral traits are similar across flying and walking. 3. The fly’s individuality remains stable across the same visual cue, independent of the locomotive modality. Subsequent quantitative analysis indeed revealed a significant correlation for angular velocity between flying and walking flies (Fig. 7b), indicating that a fly’s attention towards a visual object is, to some degree, preserved independent of the locomotor modality. In contrast, we found no statistically significant correlations for any of the exploration parameters between flying and walking. Finally, a more detailed analysis revealed mostly comparable results for male and female data (supplemental Fig. S12 and S13).

Visual attention persists across different behavioral states.

Earlier version: a Experimental overview. Flies were untethered after flying within the flight simulator for 10 min (left) and walking for 15 min (right) under the same Buridan stimulus (stripe size=11°). Single fly data (middle) and population data (bottom) reveal stripe fixation behavior when walking and during flight. b Correlation of key behavioral parameters depending on the behavioral state (flight=x-axis, walking=y-axis). For detailed fly numbers per experimental condition, see supplemental Fig. S12.

Revised version: a Experimental overview. Flies were untethered after flying within the flight simulator for 10 min (left) and walking for 15 min (right) under the same Buridan stimulus (stripe size=11°). Single fly data (middle) and group data (bottom) reveal stripe fixation behavior when walking and during flight. b Correlation of key behavioral parameters depending on the behavioral state (flight=x-axis, walking=y-axis, n=73).

Removed in the revised version: We expected one of three possible outcomes: 1. There is no stability of individual behavioral traits across flying and walking, 2. Only some behavioral traits are similar across flying and walking. 3. The fly’s individuality remains stable across the same visual cue, independent of the locomotive modality. Subsequent quantitative analysis indeed revealed a significant correlation for angular velocity between flying and walking flies (Fig. 7b), indicating that a fly’s attention is, to some degree, preserved independent of the locomotor modality. In contrast, we found no statistically significant correlations for any of the exploration parameters between flying and walking. Finally, a more detailed analysis revealed mostly comparable results for male and female data (supplemental Fig. S12 and S13).

Earlier version: In summary, the data presented here reveal a hierarchical influence of different parameters within the environmental context on individuality (Fig. 8): First, we conclude that time (at least within the frame of days) has a negligible influence on individual behavioral responses (the basis of the definition of animal individuality in a single situation), since we found stability of individual behavioral traits and parameters over time in both locomotor modalities, walking and flying. Second, temperature had the second weakest impact on individuality (although only tested for walking behavior). This contrasts the dramatically altered population mean responses, a difference that can be explained by scaled individual responses in which individuals keep the rank in the population. Third, the nature of the visual input appears to influence the individual orientation responses but has little impact on the exploration parameters. Finally, we identified the most critical situational environmental features that determine individual responses as both the arena type and the locomotor modality (walking vs. flying). Therefore, we conclude that the individuality of visually guided behaviors persists specifically in some traits across situations, but the environmental context plays a prominent role in defining individual behavioral responses.

Revised version: The data presented here reveal a hierarchical influence of different parameters within the environmental context on individuality (Fig. 8): First, we conclude that time (at least within the frame of days) has a negligible influence on individual behavioral responses (the basis of the definition of animal individuality in a single situation), since we found consistency of individual behavioral traits and parameters over time in both locomotor modalities, walking and flying. Second, temperature had the second weakest impact on individuality (although only tested for walking behavior). This contrasts the dramatically altered group mean responses, a difference that can be explained by scaled individual responses in which individuals keep the rank in the group. Third, the nature of the visual input appears to influence the individual orientation responses but has little impact on the exploration parameters. Finally, we identified the most critical situational environmental features that determine individual responses as both the arena type and the locomotor modality (walking vs. flying). Therefore, we conclude that the individuality of visually guided behaviors persists specifically in some traits across situations, but the environmental context plays a prominent role in defining individual behavioral responses.

A hierarchy of environmental and behavioral contexts based on their influence on individuality.

a Overview summarizing correlation coefficients (p<0.05) for key behavioral parameters depending on behavioral context, time, temperature, and visual input. b Quantification of the influence of behavioral and environmental contexts on the persistence of individuality.

Added in the revised version: To verify the results of our multiple correlative analyses, we took advantage of a generalized linear model (GLM, Fig. 9). We first analyzed the effect of time, temperature, vision, and arena on the mean group responses analogous to supplemental Fig. S2, S6, S10, S12, (Fig. 9a). Our reanalysis of the walking data confirmed the temperature dependency of activity parameters (Fig. S6). It also revealed effects on the group means by vision and arena particularly for walking speed. The orientation parameter vector strength was similarly affected by temperature vision and arena (see also supplemental Fig. S2, S6), while the position parameter centrophobicity was strongest affected by the arena (see also supplemental Fig. S2, S6). For flight parameters, the reanalysis of mean data revealed a strong effect of vision on the orientation parameter vector strength and the heading parameter median heading axial (see also supplemental Fig. S10). Our reanalysis, therefore, confirmed our previous analysis of mean group data.

Added in the revised version: A generalized linear model to demonstrate individuality across situations.

Added in the revised version: a The difference of individual behavioral parameters for walking (left) and flight (right) when changing environmental conditions (time, temperature, vision, arena). Significance: *p<0.05, **p<0.01, ***p<0.001. b The difference of rank (in %) within the tested groups (walking and flight combined) when changing environmental conditions. Smaller rank differences indicate a lower influence of the environmental parameter on the persistence of individuality. Significance: *p<0.05, **p<0.01, ***p<0.001. c Generalized linear model (ridge regression) showing the influence of changing environmental parameters on individual rank differences within the tested group (individuality) using the data in b. Significance: *p<0.05, **p<0.01, ***p<0.001. Non-significant environmental parameters are shaded in grey.

Added in the revised version: The next step of our reanalysis was to compare the rank changes of our flies across situations analogous to the correlation analysis we had done (Fig. 2, 4, 6, 7). Our reanalysis revealed significant changes in activity parameters and the orientation parameter vector strength. In these parameters, the effect on the rank of each individual increases from time over temperature, vision, and arena to the behavioral state. Confirming the data we found in Fig. 2, 4, 6, 7, and 8. Also, the lack of a significant effect of the behavioral state for angular velocity confirms our correlation analysis of Fig. 7 and 8.

Added in the revised version: In the last step of our reanalysis, we generated a GLM with ridge regression to show the effect of changing environmental situations on individual rank differences (Fig, 9c). For the activity parameters, time, temperature, and vision significantly stabilized individual phenotypical rank. In contrast, the arena and behavioral state significantly had a negative effect on rank stability. For the orientation parameter vector strength, we could fully replicate our previous findings that time had the strongest positive effect on consistency, followed by temperature. Vision had no significant impact. In contrast, the arena had a small negative effect consistency, while the behavioral state had a strong negative effect on behavioral consistency. In the second orientation parameter (angular velocity), a significant destabilizing effect was only found for the behavioral state. and for the positional parameter (centrophobicity), only vision had a significant effect on behavioral consistency.

Added in the revised version: In summary, using two independent mathematical approaches, we have established a hierarchy of different environmental situations on behavioral consistency. This data demonstrates that temporal consistency is also in animal most prevalent. Like humans also animals have significant consistency across situation, but similar to humans different situations evoke different behavioral responses masking measurable behavioral consistency.

Discussion

Earlier version: Our work presented here focused on the question of whether the combination of stable behavioral parameters ranging across different traits and therefore defining an animal’s individuality are invariable across situations or whether they are re-shaped by different environmental contexts. This question is closely related to the questions in the human ‘person-situation debate‘ 19. Until now, in the neurosciences individuality across situations has been largely ignored since, the commonly used definition of animal individuality only refers to the repeatability of behavior over time in a single environmental situation 1. The contrary is true for mainly field and observation-based sciences such as ethology and behavioral ecology that have since decades acquired observations for 28,29 or against 25,26 behavioral stability across situations, with some of the most convincing work probably found in hermit crabs 21,22, blue tits 23, and geese 24. Our work establishes animal individuality across situations under several tightly controlled laboratory conditions, revealing that animal individuality shares remarkable similarities to what has been described in human personality research. In agreement with Vansteelandt, 2004 our results are in line with the concept of the personality triad, emphasizing the interplay of individual, situation, and behavior (in humans: person, situation, and behavior)71.

Revised version: Our work presented here focused on the question of whether the combination of stable behavioral parameters ranging across different traits and therefore defining an animal’s individuality are invariable across situations or whether they are re-shaped by different environmental contexts. This question is closely related to the questions in the human ‘person-situation debate‘ (Mischel, 1968). Until now, in the neurosciences individuality across situations has been largely ignored since, the commonly used definition of animal individuality only refers to the repeatability of behavior over time in a single environmental situation (Gosling, 2008). The contrary is true for mainly field and observation-based sciences such as ethology and behavioral ecology that have since decades acquired observations for (Boissy & Bouissou, 1995; Spoolder et al., 1996) or against (Coleman & Wilson, 1998; D’Eath & Burn, 2002) behavioral consistency across situations, with some of the most convincing work probably found in hermit crabs (Briffa et al., 2008; Mowles et al., 2012), blue tits (Herborn et al., 2010), and geese (Kralj-Fišer et al., 2007). Our work establishes animal individuality across situations under several tightly controlled laboratory conditions, revealing that animal individuality shares remarkable similarities to what has been described in human personality research. In agreement with Vansteelandt, 2004 our results are in line with the concept of the personality triad, emphasizing the interplay of individual, situation, and behavior (in humans: person, situation, and behavior) (Vansteelandt & Van Mechelen, 2004).

Earlier version: Our results confirm that under the same environmental situation and stimulus, animals produce interindividually variable, but extremely stable behavioral responses when being retested 72. Therefore, stable and individually variable behaviors likely result from individual brain properties that persist over long timescales (in this paper days as shown in previous papers months 13), contrasting unstable random behavioral variations like sudden internal state variations. Hence, when presented with the same stimulus in the same situation, each individual responds variably but repetitively so, pointing towards a deterministic and hardwired basis of individual properties 10,13.

Revised version: Our results confirm that under the same environmental situation and stimulus, animals produce inter-individually variable, but extremely consistent behavioral responses when being retested (Bell et al., 2009). Therefore, consistent and individually variable behaviors likely result from individual brain properties that persist over long timescales (in this paper days as shown in previous papers months (Linneweber et al., 2020)), contrasting random behavioral variations like sudden internal state variations. Hence, when presented with the same stimulus in the same situation, each individual responds variably but repetitively so, pointing towards a deterministic and hardwired basis of individual properties (Buchanan et al., 2015; Linneweber et al., 2020).

Earlier version: The variability of the situational environmental context makes individual animal behavior in the wild much less predictable than that under laboratory conditions 73. In opposition to the vast human literature on the ‘person-situation debate’ 19,31,33,71,74, the importance of environmental context in shaping individual animal behavior remains descriptive 24 and experimentally largely unexplored, with few exceptions 21,22,27,75. Here we closed this gap in knowledge by quantitatively demonstrating how situational environmental and locomotor context changes individual behavior and individuality in animals. Our results demonstrate that situational environmental context has a pronounced effect on individual behavioral responses, be it an alternative stimulus representation, a different arena, or another locomotive modality.

Revised version: The variability of the situational environmental context makes individual animal behavior in the wild much less predictable than that under laboratory conditions (Briffa & Greenaway, 2011). In opposition to the vast human literature on the ‘person-situation debate’ (Diener & Larsen, 1984; Fleeson, 2001; Mischel, 1968; Vansteelandt & Van Mechelen, 2004; Weisbuch et al., 2010), the importance of environmental context in shaping individual animal behavior remains descriptive (Kralj-Fišer et al., 2007) and experimentally largely unexplored, with few exceptions (Briffa et al., 2008; Mowles et al., 2012; Versace et al., 2020; Werkhoven et al., 2021). Here, we closed this gap in knowledge by quantitatively demonstrating how situational environmental and locomotor context changes individual behavior and individuality in animals. Our results demonstrate that situational environmental context has a pronounced effect on individual behavioral responses, be it an alternative stimulus representation, a different arena, or another locomotive modality.

Earlier version: Similar to the human literature 19, the correlation coefficients are lower across situations than in the same situation or environmental context. From this, we can conclude that under more natural conditions with varying contexts, the stability of individual traits will be reduced but remain measurable; therefore, for each situational context, the brain computes the appropriate behavior. This is like the situation-behavior profile of the individual in the personality psychology literature 76,77. Each time the individual is placed in the same situational context, the same behavior will be evoked, in a different context, a modified behavior. But despite these modifications, the individual remains stable. In line with this, a vast literature demonstrates the importance of animal individual differences for ecological communities under natural conditions reviewed in: 20,78.

Revised version: Similar to the human literature (Mischel, 1968), the correlation coefficients are lower across situations than in the same situation or environmental context. From this, we can conclude that under more natural conditions with varying contexts, the consistency of individual traits measured by behavioral parameters will be reduced but remain measurable; therefore, for each situational context, the brain computes the appropriate behavior. This is like the situation-behavior profile of the individual in the personality psychology literature (Mischel & Shoda, 1995; Mischel & Shoda, 1998). Each time the individual is placed in the same situational context, the same behavior will be evoked, in a different context, a modified behavior. But despite these modifications, the individual brain properties computing the response remain the same. In line with lower but measurable behavioral consistency across situations, a vast literature demonstrates the importance of animal individual differences for ecological communities under natural conditions (reviewed in: Sih et al., 2012; Takola & Schielzeth, 2022).

Earlier version: Our quantitative analysis of the effect of environmental and locomotor context on behavioral stability revealed a hierarchy. As expected from animal and human studies, the time frame of days had little effect on the behavioral outcome 8,10,13,19,32,72,79. The hierarchy revealed by our experiments argues for differences in the relevance of environmental features for the traits we measured, which aligns with Mischel’s strong or weak situations 80. Therefore, it is unsurprising that exploration can be best measured across situations within experiments of walking flies, and attention towards a stimulus can be measured across movement modalities. Our quantitative analysis demonstrates that situational context can be the most significant factor in determining the individual behavioral outcome, even in a per-design visual assay. This argues that individual behavior is determined by the entirety of the situation consisting of both the stimulus and the environmental context. For each situation, each individual manifests its individual behavior.

Revised version: Our quantitative analysis of the effect of environmental and locomotor context on behavioral consistency revealed a hierarchy. As expected from animal and human studies, the time frame of days had little effect on the behavioral outcome (Bell et al., 2009; Buchanan et al., 2015; Linneweber et al., 2020; Mischel, 1968; Pantoja et al., 2016; Ross, 1977; Schuett et al., 2011). The hierarchy revealed by our experiments argues for differences in the relevance of environmental features for the traits we measured, which aligns with Mischel’s strong or weak situations (Mischel, 1973). Therefore, it is unsurprising that exploration can be best measured across situations within experiments of walking flies, and attention towards a stimulus can be measured across movement modalities. Our quantitative analysis demonstrates that situational context can be the most significant factor in determining the individual behavioral outcome, even in a per-design visual assay. This argues that individual behavior is determined by the entirety of the situation consisting of both the stimulus and the environmental context. For each situation, each individual manifests its individual behavior.

In the future, it will be important to investigate which individual neurophysiological responses change when tested under exactly the same stimulus, but different situations. A systematic analysis of relevant circuit elements for different behaviors will allow for determining the exact neuronal substrates underlying individual behavioral computations within a fly’s or other animals’ brains. This future research on animal models will provide important new insight into the ‘person-situation debate‘.

Earlier version: On a more practical level, the importance of the situational environmental context for shaping individual behavioral outcomes is also very important when comparing the results of different research groups when using similar behavioral paradigms. In many instances, it is impossible to replicate all components of a specific assay, let alone the conditions of the experimental room. Our results show that such details substantially impact individual behavioral responses and contribute to errors in replicating the results from another lab 81,82. Only meticulous and detailed reporting of as many environmental features as possible will hopefully alleviate some of these problems in the future.

Revised version: On a more practical level, the importance of the situational environmental context for shaping individual behavioral outcomes is also very important when comparing the results of different research groups when using similar behavioral paradigms. In many instances, it is impossible to replicate all components of a specific assay, let alone the conditions of the experimental room. Our results show that such details substantially impact individual behavioral responses and contribute to errors in replicating the results from another lab (Crabbe et al., 1999; Lithgow et al., 2017). Only meticulous and detailed reporting of as many environmental features as possible will hopefully alleviate some of these problems in the future. In conclusion, our work confirms that animals including flies have, like humans, individuality across situations. Similar to humans, animal individuality across situations is also marked by lower correlation coefficients than repetitions of a single situation. The individual is stable across situations producing individualistic behavior through an interaction with the situation. This involves a hierarchy of influence of different environmental features on the predictability of behavior across situations.

Removed in the revised version: In conclusion, our work confirms that animals including flies have, like humans, individuality across situations. Similar to humans, animal individuality across situations is also marked by lower correlation coefficients than repetitions of a single situation. The individual is stable across situations producing individualistic behavior through an interaction with the situation. This involves a hierarchy of influence of different environmental features on the predictability of behavior across situations.